# 1. feladatsor - Magasabb rendű függvények

## Oszthatóság számolása

Készíts egy Python függvény gyártó függvényt ami képes futás közben olyan függvényeket előállítani melyek megadott számmal való oszthatóságot képesek megállapítani. 

PL:
```python
is_div_2 = make_is_div(2)
assert(is_div_2(234) == True)
```

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
def make_is_div(n):
    def is_div(x):
        return x % n == 0
    return is_div

### END SOLUTION
is_div_2 = make_is_div(2)
print(f'oszthato 234 2-vel? {is_div_2(234)}')
```

Output:
```
oszthato 234 2-vel? True
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_true("make_is_div" in locals())
div_2 = make_is_div(2)
assert_true(div_2(0) == True)
assert_true(div_2(1) == False)
assert_true(div_2(2) == True)
assert_true(div_2(234) == True)

div_5 = make_is_div(5)
assert_true(div_5(0) == True)
assert_true(div_5(1) == False)
assert_true(div_5(2) == False)
assert_true(div_5(5) == True)
assert_true(div_5(10) == True)
assert_true(div_5(234) == False)
```

## Oszthatóság 10000ig
Készítsd el az első 10000 számmal való oszthatóságot megállapító függvényeket a `divs` változóba és írd ki az eredményeket a `numbers_to_test` példa számokkal, pl:

112233 szamra futtatva a kiírandó:
`[1, 3, 11, 19, 33, 57, 179, 209, 537, 627, 1969, 3401, 5907]`

<!-- grade_id: cell-f6ebbd492f6702c0, points: 0, grade: false, solution: true -->
```python
divs = []
numbers_to_test = [2343453, 1337, 666777, 5552213]
## BEGIN SOLUTION
divs = [make_is_div(i) for i in range(1, 10001)]
for num in numbers_to_test:
    print([i+1 for i, is_div in enumerate(divs) if is_div(num)])
### END SOLUTION
```

Output:
```
[1, 3, 7, 21]
[1, 7, 191, 1337]
[1, 3, 37, 111, 6007]
[1, 1061, 5233]
```

<!-- grade_id: cell-f7bf2ba072825047, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution) == 4)
assert_true(solution[0] == '[1, 3, 7, 21]')
assert_true(solution[1] == '[1, 7, 191, 1337]')
assert_true(solution[2] == '[1, 3, 37, 111, 6007]')
assert_true(solution[3] == '[1, 1061, 5233]')

```

## Saját apply függvény

Készíts egy saját magasabb rendű függvényt. A függvény várjon el a bemenetén egy string listát és egy függvényt. A függvény nyomtassa ki az összes elemet a listából miután meghívta a függvényt az adott elemre. Legyen a neve `my_apply`

Pl:
```python
my_apply(lambda s: s.upper(), ["Hello, World", "Tamas", "big", "world"])
=>
HELLO, WORLD
TAMAS
BIG
WORLD
```

<!-- grade_id: cell-1c77d962245935e8, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
def my_apply(fn, lst):
    for i in lst:
        print(fn(i))

### END SOLUTION
my_apply(lambda s: s.upper(), ["Hello, World", "Tamas", "big", "world"])
```

Output:
```
HELLO, WORLD
TAMAS
BIG
WORLD
```

<!-- grade_id: cell-a0b6590684e1355c, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true('my_apply' in locals())

assert_true(len(solution) == 4)
assert_true(solution[0] == 'HELLO, WORLD')
assert_true(solution[1] == 'TAMAS')
assert_true(solution[2] == 'BIG')
assert_true(solution[3] == 'WORLD')
```

# 2. feladatsor - generátorok

## Egyszerű generátor
Írj egy Python generátor kifejezést (generator expression) ami kiszámítja a köbét minden páros számnak az alábbi listából `[10,3,7,9,11,33,22,8,2]`.

Használd a `list` függvényt, hogy az eredményeket egy listában tudd prezentálni.

<!-- grade_id: cell-780ded4473d0d05f, points: 0, grade: false, solution: true -->
```python
test_list =  [10, 3, 7, 9, 11, 33, 22, 8, 2]

## BEGIN SOLUTION
print(list(x ** 3 for x in test_list if x % 2 == 0))
### END SOLUTION
```

Output:
```
[1000, 10648, 512, 8]
```

<!-- grade_id: cell-5a4a6eda3b12bf3b, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution) == 1)
assert_true(solution[0] == '[1000, 10648, 512, 8]')
```

## Generátor pipeline

Készítsünk Python generátor pipline-t, azaz ágyazzunk egymásba több generátor kifejezést (generator expression).


a) Írj egy python generátor kifejezést ami a `numbers` listából csinál generátort.


b) Írj egy másik kifejezést ami négyzetre négyzetre emeli a számokat


c) Írj egy harmadik kijezést ami stringge konvertálja az értékeket


d) Ellenőrzés képpen irasd ki az összes egyedi stringet amit kaptunk. (Hasznos lehet a `set` halmaz adatszerkezet)

<!-- grade_id: cell-f5cef025ec94281e, points: 0, grade: false, solution: true -->
```python
numbers = [80, 251, 653, 541, 169, 949, 114, 35, 53, 348, 799, 681, 595, 273, 151, 797, 42, 171, 477, 323, 463, 208, 736, 859, 49, 594, 827, 4, 776, 308, 703, 451, 128, 66, 386, 875, 344, 322, 903, 787, 294, 177, 690, 24, 182, 589, 122, 968, 692, 472, 875, 973, 579, 168, 945, 388, 19, 228, 960, 569, 909, 297, 103, 477, 184, 835, 483, 547, 487, 762, 422, 326, 897, 451, 772, 101, 960, 125, 608, 216, 29, 556, 286, 295, 833, 634, 297, 296, 545, 494, 828, 58, 407, 874, 709, 294, 590, 888, 183, 51]
## BEGIN SOLUTION
randomizer = (n for n in numbers)
squared = (x ** 2 for x in randomizer)
stringified = (str(x) for x in squared)
print(list(set(list(stringified))))
### END SOLUTION
```

Output:
```
['602176', '88209', '214369', '841', '2601', '946729', '148996', '580644', '635209', '10609', '697225', '369664', '15625', '74529', '638401', '323761', '28561', '14884', '937024', '541696', '6400', '478864', '31329', '826281', '893025', '619369', '16384', '63001', '292681', '299209', '502681', '788544', '1225', '685584', '121104', '165649', '494209', '33489', '46656', '222784', '12996', '244036', '576', '463761', '233289', '352836', '763876', '2401', '86436', '118336', '309136', '81796', '227529', '87616', '203401', '106276', '683929', '103684', '335241', '178084', '354025', '10201', '900601', '346921', '815409', '737881', '237169', '804609', '28224', '2809', '33856', '22801', '104329', '43264', '693889', '476100', '1764', '297025', '426409', '3364', '595984', '29241', '765625', '150544', '921600', '87025', '401956', '348100', '16', '51984', '4356', '361', '33124', '94864']
```

<!-- grade_id: cell-1e746922130068c9, points: 1, grade: true, solution: false -->
```python
import ast 

solution = %solution
assert_true(len(solution) == 1)
assert_equal(set(ast.literal_eval(solution[0])), set(['602176', '88209', '214369', '841', '2601', '946729', '148996', '580644', '635209', '10609', '697225', '369664', '15625', '74529', '638401', '323761', '28561', '14884', '937024', '541696', '6400', '478864', '31329', '826281', '893025', '619369', '16384', '63001', '292681', '299209', '502681', '788544', '1225', '685584', '121104', '165649', '494209', '33489', '46656', '222784', '12996', '244036', '576', '463761', '233289', '352836', '763876', '2401', '86436', '118336', '309136', '81796', '227529', '87616', '203401', '106276', '683929', '103684', '335241', '178084', '354025', '10201', '900601', '346921', '815409', '737881', '237169', '804609', '28224', '2809', '33856', '22801', '104329', '43264', '693889', '476100', '1764', '297025', '426409', '3364', '595984', '29241', '765625', '150544', '921600', '87025', '401956', '348100', '16', '51984', '4356', '361', '33124', '94864']))
```

## Egyedi elemek
írj egy Python generátor függvényt `unique` néven ami kiválasztja az egyedi elemeket bármilyen listából.

<!-- grade_id: cell-cc46d21d001284ad, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
def unique(iterable):
    for i in set(iterable):
        yield i
    
### END SOLUTION
u = unique([1,'Tamas', 2,1,5,32,2,1, 'Agi', 33.3, 33.3, 33.3])

for i in u:
   print(i) 
```

Output:
```
32
1
2
33.3
5
Tamas
Agi
```

<!-- grade_id: cell-da3a04638c5a70b5, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution) == 7)
assert_equal(set(solution), set(['32', '1', '2', '33.3', '5', 'Tamas', 'Agi']))
```

# 3. feladatsor - itertools
## Töltés jelző

Készíts egy Python programot ami egy parancssori progress indikátort fog megjeleníteni 5 másodpercig, 0.1 másodpercenként frissítve a kiírást.


Az indikátor legyen karakteres ("\\|/-"). Használd az itertools könyvtárat. (segítség, a cycle() függvény jól jöhet, illetve az is, hogy a print függvénynél a end='\r' opcionális paraméter visszatér a sor elejére).

Példa lépések:
```python
\
|
/
-
\
...
```



<!-- grade_id: cell-902e5eae6cb9c883, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
import itertools as it
import time

ind_chars = "\|/-"

indicator = it.cycle(ind_chars)

for _ in range(100):
    print(next(indicator), end='\r')
    time.sleep(0.1)
### END SOLUTION
```

Output:
```
-
```

<!-- grade_id: cell-7b29fdd26a08cc07, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_equal(solution,['\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-', '\\', '|', '/', '-'])
```

## Csoportba rendezés

Készíts Python függvényt `shift_maker` néven, ami munkahelyi "váltásokba" rendez embereket. Egy váltás például `('Peter, 'Julia')`. A visszatérési érték legyen egy lista ami tartalmazza a generált csoportokat.

Bónusz feladat: randomizáld a kiválasztást.

```python
shift_maker(people_iterable, number_of_shifts)

shift_maker(['Peter', 'Julia', 'Andrew', 'Bobby', 'Margo', 'Bill', 'Alice', 'Anna'], 3)
```

<!-- grade_id: cell-fc2db0891925c813, points: 0, grade: false, solution: true -->
```python
people = ['Peter', 'Julia', 'Andrew', 'Bobby', 'Margo', 'Bill', 'Alice', 'Anna']
## BEGIN SOLUTION
import itertools as it
def shift_maker(people_iterable, num_of_shifts):
    people_per_shift = len(people_iterable) / num_of_shifts
    return it.combinations(people_iterable, int(people_per_shift))

### END SOLUTION
print(list(shift_maker(people, 3)))
```

Output:
```
[('Peter', 'Julia'), ('Peter', 'Andrew'), ('Peter', 'Bobby'), ('Peter', 'Margo'), ('Peter', 'Bill'), ('Peter', 'Alice'), ('Peter', 'Anna'), ('Julia', 'Andrew'), ('Julia', 'Bobby'), ('Julia', 'Margo'), ('Julia', 'Bill'), ('Julia', 'Alice'), ('Julia', 'Anna'), ('Andrew', 'Bobby'), ('Andrew', 'Margo'), ('Andrew', 'Bill'), ('Andrew', 'Alice'), ('Andrew', 'Anna'), ('Bobby', 'Margo'), ('Bobby', 'Bill'), ('Bobby', 'Alice'), ('Bobby', 'Anna'), ('Margo', 'Bill'), ('Margo', 'Alice'), ('Margo', 'Anna'), ('Bill', 'Alice'), ('Bill', 'Anna'), ('Alice', 'Anna')]
```

<!-- grade_id: cell-ce7c1fa067a5d3e3, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(list(shift_maker(people, 2)), [('Peter', 'Julia', 'Andrew', 'Bobby'), ('Peter', 'Julia', 'Andrew', 'Margo'), ('Peter', 'Julia', 'Andrew', 'Bill'), ('Peter', 'Julia', 'Andrew', 'Alice'), ('Peter', 'Julia', 'Andrew', 'Anna'), ('Peter', 'Julia', 'Bobby', 'Margo'), ('Peter', 'Julia', 'Bobby', 'Bill'), ('Peter', 'Julia', 'Bobby', 'Alice'), ('Peter', 'Julia', 'Bobby', 'Anna'), ('Peter', 'Julia', 'Margo', 'Bill'), ('Peter', 'Julia', 'Margo', 'Alice'), ('Peter', 'Julia', 'Margo', 'Anna'), ('Peter', 'Julia', 'Bill', 'Alice'), ('Peter', 'Julia', 'Bill', 'Anna'), ('Peter', 'Julia', 'Alice', 'Anna'), ('Peter', 'Andrew', 'Bobby', 'Margo'), ('Peter', 'Andrew', 'Bobby', 'Bill'), ('Peter', 'Andrew', 'Bobby', 'Alice'), ('Peter', 'Andrew', 'Bobby', 'Anna'), ('Peter', 'Andrew', 'Margo', 'Bill'), ('Peter', 'Andrew', 'Margo', 'Alice'), ('Peter', 'Andrew', 'Margo', 'Anna'), ('Peter', 'Andrew', 'Bill', 'Alice'), ('Peter', 'Andrew', 'Bill', 'Anna'), ('Peter', 'Andrew', 'Alice', 'Anna'), ('Peter', 'Bobby', 'Margo', 'Bill'), ('Peter', 'Bobby', 'Margo', 'Alice'), ('Peter', 'Bobby', 'Margo', 'Anna'), ('Peter', 'Bobby', 'Bill', 'Alice'), ('Peter', 'Bobby', 'Bill', 'Anna'), ('Peter', 'Bobby', 'Alice', 'Anna'), ('Peter', 'Margo', 'Bill', 'Alice'), ('Peter', 'Margo', 'Bill', 'Anna'), ('Peter', 'Margo', 'Alice', 'Anna'), ('Peter', 'Bill', 'Alice', 'Anna'), ('Julia', 'Andrew', 'Bobby', 'Margo'), ('Julia', 'Andrew', 'Bobby', 'Bill'), ('Julia', 'Andrew', 'Bobby', 'Alice'), ('Julia', 'Andrew', 'Bobby', 'Anna'), ('Julia', 'Andrew', 'Margo', 'Bill'), ('Julia', 'Andrew', 'Margo', 'Alice'), ('Julia', 'Andrew', 'Margo', 'Anna'), ('Julia', 'Andrew', 'Bill', 'Alice'), ('Julia', 'Andrew', 'Bill', 'Anna'), ('Julia', 'Andrew', 'Alice', 'Anna'), ('Julia', 'Bobby', 'Margo', 'Bill'), ('Julia', 'Bobby', 'Margo', 'Alice'), ('Julia', 'Bobby', 'Margo', 'Anna'), ('Julia', 'Bobby', 'Bill', 'Alice'), ('Julia', 'Bobby', 'Bill', 'Anna'), ('Julia', 'Bobby', 'Alice', 'Anna'), ('Julia', 'Margo', 'Bill', 'Alice'), ('Julia', 'Margo', 'Bill', 'Anna'), ('Julia', 'Margo', 'Alice', 'Anna'), ('Julia', 'Bill', 'Alice', 'Anna'), ('Andrew', 'Bobby', 'Margo', 'Bill'), ('Andrew', 'Bobby', 'Margo', 'Alice'), ('Andrew', 'Bobby', 'Margo', 'Anna'), ('Andrew', 'Bobby', 'Bill', 'Alice'), ('Andrew', 'Bobby', 'Bill', 'Anna'), ('Andrew', 'Bobby', 'Alice', 'Anna'), ('Andrew', 'Margo', 'Bill', 'Alice'), ('Andrew', 'Margo', 'Bill', 'Anna'), ('Andrew', 'Margo', 'Alice', 'Anna'), ('Andrew', 'Bill', 'Alice', 'Anna'), ('Bobby', 'Margo', 'Bill', 'Alice'), ('Bobby', 'Margo', 'Bill', 'Anna'), ('Bobby', 'Margo', 'Alice', 'Anna'), ('Bobby', 'Bill', 'Alice', 'Anna'), ('Margo', 'Bill', 'Alice', 'Anna')]
)
assert_equal(list(shift_maker(people, 3)), [('Peter', 'Julia'), ('Peter', 'Andrew'), ('Peter', 'Bobby'), ('Peter', 'Margo'), ('Peter', 'Bill'), ('Peter', 'Alice'), ('Peter', 'Anna'), ('Julia', 'Andrew'), ('Julia', 'Bobby'), ('Julia', 'Margo'), ('Julia', 'Bill'), ('Julia', 'Alice'), ('Julia', 'Anna'), ('Andrew', 'Bobby'), ('Andrew', 'Margo'), ('Andrew', 'Bill'), ('Andrew', 'Alice'), ('Andrew', 'Anna'), ('Bobby', 'Margo'), ('Bobby', 'Bill'), ('Bobby', 'Alice'), ('Bobby', 'Anna'), ('Margo', 'Bill'), ('Margo', 'Alice'), ('Margo', 'Anna'), ('Bill', 'Alice'), ('Bill', 'Anna'), ('Alice', 'Anna')])

```

## Üzenetek
A hálózaton számokat fogadunk, amik egy listába kerülnek letárolásra. A nem 0 karakterek egymás után egy blokkot alkotnak, a 0 karakterek pedig "csendet" jelentenek.

Írj egy Python programot, ami az eredeti listából kiválogatja az értelmezhető blokkokat, és előállít egy olyan listát amikben csak ezek vannak benne.

Pl:
```python
[1,3,5,0,0,0,0,7,7,0,0,11,0,5,6,7] => [[1,3,5], [7,7], [11], [5,6,7]]
```


<!-- grade_id: cell-aba5370efa4a5304, points: 0, grade: false, solution: true -->
```python
import itertools as it
nums = [3,4,6,2,0,0,0,0,0,0,6,7,6,9,10,0,0,0,0,0,7,4,4,0,0,0,0,0,0,5,3,2,9,7,1]

## BEGIN SOLUTION
def solve(lst):
    result = [ list(x[1]) for x in it.groupby(lst, lambda x: x == 0) if not x[0] ]
    return result
print(solve(nums))
### END SOLUTION
```

Output:
```
[[3, 4, 6, 2], [6, 7, 6, 9, 10], [7, 4, 4], [5, 3, 2, 9, 7, 1]]
```

<!-- grade_id: cell-9e43725030452de0, points: 1, grade: true, solution: false -->
```python
solution = %solution
import ast
assert_equal(ast.literal_eval(solution[0]), [[3, 4, 6, 2], [6, 7, 6, 9, 10], [7, 4, 4], [5, 3, 2, 9, 7, 1]])
```