# 1. feladatsor - Dekorátorok

## Oszthatóság számolása

Készíts egy Python dekorátort `timeit_decorator` néven ami képes lemérni es printelni egy függvény időkomplexitását, azaz azt az időt ami a futtatásához kellett. 

PL:

```
@timeit_decorator
def comprehension():

    return [x**2 for x in range(3000000)]
```

A kiírás elvárt formája:

`FUNCTION_NAME took TIME`


`comprehension took 767.8432464599609`


<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
import time

## BEGIN SOLUTION
def timeit_decorator(fn):
    def wrapper(*args):
        start = time.time()
        res = fn(*args)
        end = time.time()
        print(f'{fn.__name__} took {(end - start)*1000}')
        return res
    return wrapper
### END SOLUTION

@timeit_decorator
def comprehension():
    return [x**2 for x in range(3000000)]

@timeit_decorator
def generator():
    return (x**2 for x in range(3000000))

@timeit_decorator
def iter_over(iter):
    for _ in iter:
        pass

    
alist = comprehension()
iter_over(alist)
agenerator = generator()
iter_over(agenerator)
```

Output:
```
comprehension took 679.8884868621826
iter_over took 33.86116027832031
generator took 0.0021457672119140625
iter_over took 736.0579967498779
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(len(solution), 4)
assert_true("comprehension took" in solution[0])
assert_true("iter_over took" in solution[1])
assert_true("generator took" in solution[2])
assert_true("iter_over took" in solution[3])
```

## Anonimizáló dekorátor
Készíts egy dekorátort `anonimize` néven, ami egy függvény visszatérési értékét anonimizálja, azaz olvashatatlanná teszi, ha a visszatérési érték `str` típusú.
Tegye ezt úgy, hogy minden karaktert `#`-é alakít.
Pl:
    
`almafa -> ######`




<!-- grade_id: cell-5c227f0b8eec859e, points: 0, grade: false, solution: true -->
```python
import random
## BEGIN SOLUTION
def anonimize(fn):
    def wrapper(*args):
        res = fn(*args)
        if type(res) == str:
            return "#" * len(res)
        return res
    return wrapper
### END SOLUTION

@anonimize
def get_result():
    data = ["alma", "kortefa", "barack", "bananfa"]
    return random.choice(data)


print(get_result())
```

Output:
```
######
```

<!-- grade_id: cell-6658e55c4a75eb0e, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true('#' in get_result())
assert_true('#' in get_result())
assert_true('#' in get_result())
assert_true('#' in get_result())
assert_true('#' in get_result())
```

## Logging dekorátor

Írj egy olyan dekorátort `logged` néven, ami egy változó hosszú paraméter listával rendelkező függvényt be tud csomagolni és ki tudja írni a függvény összes paraméterét, nevét, valamint azt, hogy hány paramétere van.

Kiírási minta:

```python
hello(2,4,6)
=>
you called hello((2, 4, 6))
result was 6
```

<!-- grade_id: cell-ab64cbd1a8753849, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
import inspect

def logged(fn):
    def wrapper(*args):
        result = fn(*args)
        print(f'you called {fn.__name__}({args})')
        print(f'result was {result}')
        return result
    return wrapper
### END SOLUTION

@logged
def hello(*args):
    return 3 + len(args)

hello(2,4,6)
```

Output:
```
you called hello((2, 4, 6))
result was 6
6
```

<!-- grade_id: cell-be785a9ad12141b5, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(solution[0], "you called hello((2, 4, 6))")
assert_equal(solution[1], "result was 6")

```

## 2. feladatsor - Haladó objektum orientált programozás

Készíts egy Python programot ami az alábbi osztályokat és interface-eket valósítja meg:

```ascii
    ┌─────────────────────────┐            ┌─────────────────────────┐
    │       <<abc.ABC>>       │            │                         │
    │          Bird           │            │         Parrot          │
    ├─────────────────────────┤            ├─────────────────────────┤
    │                         │◄───────────┤                         │
    ├─────────────────────────┤            ├─────────────────────────┤
    │         *fly()          │            │         *fly()          │
    └─────────────────────────┘            └─────────────────────────┘
    ┌─────────────────────────┐            ┌─────────────────────────┐
    │       <<abc.ABC>>       │            │                         │
    │          Hero           │            │        SuperMan         │
    ├─────────────────────────┤            ├─────────────────────────┤
    │                         │◄───────────┤                         │
    ├─────────────────────────┤            ├─────────────────────────┤
    │         *fly()          │            │         *fly()          │
    ├─────────────────────────┤            ├─────────────────────────┤
    │       *use_power()      │            │       *use_power()      │
    └─────────────────────────┘            └─────────────────────────┘
```

Ügyelj rá, hogy a formális interface-t használj.

A példa futása az alábbi kimenetet kell adnia:
```text
slow flyer
superpower
False
False
```

<!-- grade_id: cell-959a7afa3ad2a3fa, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
import abc


class Bird(abc.ABC):
    @abc.abstractmethod
    def fly(self):
        pass


class Hero(abc.ABC):
    @abc.abstractmethod
    def use_power(self):
        pass

    @abc.abstractmethod
    def fly(self):
        pass


class Parrot(Bird):
    def fly(self):
        print('slow flyer')


class SuperMan(Hero):
    def use_power(self):
        print('superpower')

    def fly(self):
        print("It's not a bird, it is superman")
### END SOLUTION


p = Parrot()

p.fly()

super_man = SuperMan()

super_man.use_power()

print(isinstance(p, Hero))

print(isinstance(super_man, Bird))
```

Output:
```
slow flyer
superpower
False
False
```

<!-- grade_id: cell-d63bdf5f3da38857, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_false(isinstance(p, Hero))
assert_false(isinstance(super_man, Bird))
assert_true(isinstance(p, Bird))
assert_true(isinstance(super_man, Hero))
assert_true(hasattr(super_man, 'use_power'))
assert_true(hasattr(super_man, 'fly'))
assert_true(hasattr(p, 'fly'))
assert_false(hasattr(p, 'use_power'))
```

## 3. feladatsor - többszörös öröklődés
Készíts el egy Python programot ami az alábbi objektum hierarchiát és működést valósítja meg:

```ascii
                         +---------------------------+
                         |          Address          |    +---------------------------+
                         +---------------------------+    |          Person           |
                         |          street           |    +---------------------------+
                         |           city            |    |                           |
                         +---------------------------+    +---------------------------+
                         |__init__(street, city)     |    |__init__(name, email)      |
                         |show()                     |    |show()                     |
                         +-------------+-------------+    +-------------------+-------+
                                       |                                      |
                                       |                                      |
                                       |                                      |
                                       |                                      |
                                       |                                      |
+---------------------------+          +>+-----------------------------+<-----+
|         Notebook          |            |         Contact             |
+------------------------+--+  1   0..1  +-----------------------------+
|    people:dict(Contact)|+++----------->|                             |
+------------------------+--+            +-----------------------------+
|add(name, email, street,   |            |__init__(name, email, street,|
|    city)                  |            |         city)               |
|show(name)                 |            |show()                       |
+---------------------------+            +-----------------------------+
```

A többes öröklődés a Contact, Address és Person objektumok között áll fenn. A teszteléshez egy Notebook osztályt készíts ami használja a Contact osztályt. Tárold a ”people” nevű példány változóban a Contact-otkat. Ez legyen egy dict. A show metódus a Notebookon legyen képes megjeleníteni az egyes Contact-okat név alapján.

Elvárt kimenet:

```text
Alice <al@kth.se>
Lv 24
Sthlm

Unknown Carol
```

<!-- grade_id: cell-8212d83c0e0ccc47, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
class Address:
    def __init__(self, street, city):
        self.street = str(street)
        self.city = str(city)
    def show(self):
        print(self.street)
        print(self.city)

class Person:
    def __init__(self, name, email):
        self.name = str(name)
        self.email= str(email)
    def show(self):
        print(self.name + ' ' + self.email)

class Contact(Person, Address):
    def __init__(self, name, email, street, city):
        Person.__init__(self, name, email)
        Address.__init__(self, street, city)
    def show(self):
        Person.show(self)
        Address.show(self)
        print()

class Notebook:
    people = dict()
    def add(self, name, email, street, city):
        self.people[name] = Contact(name, email, street, city)
    def show(self, name):
        if name in self.people:
            self.people[name].show()
        else:
            print('Unknown', name)

### END SOLUTION
notes = Notebook()
notes.add('Alice', '<al@kth.se>', 'Lv 24', 'Sthlm')
notes.add('Bob', '<bb@kth.se>', 'Rtb 35', 'Sthlm')

notes.show('Alice')
notes.show('Carol')
```

Output:
```
Alice <al@kth.se>
Lv 24
Sthlm

Unknown Carol
{'Alice': <__main__.Contact object at 0x7faa22be5880>, 'Bob': <__main__.Contact object at 0x7faa22be5370>}
```

<!-- grade_id: cell-ab973b6234b288bc, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(type(notes.people) == dict)
assert_true(type(notes.people['Alice']) == Contact)
assert_true(isinstance(notes.people['Alice'],Person))
assert_true(isinstance(notes.people['Alice'],Address))
```