# 1. feladatsor - Többszálú programozás

Heads up: Előfordulhat, hogy a validáció kiakad a multiprocessing miatt. Többször újra futtatva egyszer csak jó lesz (5ből 1x). Ha nem működik, akkor próbálj meg saját környezetben fejleszteni.

## Összegek és szorzatok

Készíts Python függvényt `add_mul` néven, ami kiszámítja két szám összegét és szorzatát , majd kiírja ezeket. A függvényt futtasd le párhuzamosan az `in_numbers` egész számpárokat tartamazó listára. Használd a threading modult.

Tárold el az eredményt, és írj ki egy listát amiben az eredménypárok vannak (az összes threadről)

A kiírás formátuma:
```text
[(112679, 2020619928), (109787, 1110548322), ... ]
```



<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
import threading
import queue
import random
import time

in_numbers = [(90303, 22376), (11273, 98514), (64060, 33635), (22089, 8050), (72660, 31099), (23101, 54786), (31038, 7788), (25814, 56676), (10704, 44539), (56234, 17328), (4461, 4011), (67005, 71346), (34680, 75134), (39880, 93287), (96237, 39684), (75533, 4581), (64673, 80677), (88697, 23255), (58604, 2588), (93777, 27296), (68452, 47199), (8752, 58413), (81743, 42804), (28604, 25564), (90382, 72751), (94608, 69846), (99486, 55805), (31045, 79984), (23501, 1292), (97615, 96230), (7981, 13814), (64278, 89607), (14369, 16458), (91281, 13725), (60680, 83198), (7918, 86438), (75796, 48039), (30980, 16266), (4783, 65034), (63216, 7524), (75564, 67108), (6343, 34765), (96831, 41889), (17360, 19283), (92437, 38028), (72008, 24110), (62439, 82031), (13640, 67463), (29424, 88371), (50037, 30068), (91456, 47154), (69640, 59107), (64092, 74423), (52108, 66132), (21016, 34532), (77334, 90246), (9885, 34343), (81428, 84097), (40454, 1363), (15523, 38709), (31890, 12290), (22936, 80193), (66163, 41697), (85549, 15690), (88271, 22227), (98651, 36265), (87749, 68807), (41933, 12269), (32073, 75413), (62073, 75309), (24130, 2631), (61296, 50833), (26630, 33759), (70541, 39226), (1854, 49280), (9035, 52003), (41170, 46486), (52176, 6544), (12616, 93746), (71654, 99024), (52576, 64557), (84850, 53706), (1747, 11507), (39891, 33685), (88925, 99106), (84162, 67553), (49430, 76564), (90610, 59695), (89416, 37500), (18693, 88264), (20623, 2398), (64220, 19929), (94855, 44282), (81873, 19837), (41495, 78952), (72962, 98651), (56916, 78590), (8941, 84621), (3690, 11051), (52925, 96578), (47519, 82630), (59781, 21353), (92598, 88914), (47276, 65272), (49924, 45597), (38566, 39245), (47442, 79582), (99900, 31851), (53408, 25951), (98448, 19179), (7905, 90655), (33360, 15684), (2542, 73938), (83991, 68416), (47845, 37567), (4018, 80565), (78621, 88513), (69613, 30745), (43067, 74874), (48425, 69715), (89708, 85572), (77058, 15513), (13430, 87594), (17974, 79921), (50883, 32133), (93995, 27855), (20902, 12152), (11838, 7788), (43259, 19978), (43449, 6423), (75213, 25329), (85927, 55875), (72254, 25244), (60634, 14239), (52867, 48001), (36559, 28213), (61387, 25879), (81579, 97856), (83025, 15183), (54787, 33365), (77432, 85827), (2565, 3297), (82022, 80935), (32549, 4822), (7959, 578), (76859, 94457), (14671, 63755), (50976, 31016), (89383, 49413), (41710, 81302), (78348, 95131), (75920, 50653), (78533, 32735), (76695, 35263), (68611, 48007), (20824, 12614), (18330, 32697), (54150, 95769), (37603, 14346), (66287, 20627), (18237, 65292), (58492, 23727), (50244, 3577), (29398, 71178), (87772, 85878), (51556, 27639), (25796, 93509), (99769, 75886), (30627, 34156), (68592, 85475), (48315, 56384), (66320, 58504), (48654, 86183), (65022, 39085), (94555, 38293), (35913, 64595), (24590, 49779), (40005, 48343), (74945, 29140), (71823, 52363), (33931, 28266), (61702, 31460), (35340, 2004), (71330, 59352), (43502, 28095), (49934, 12122), (4453, 53916), (91115, 8495), (84990, 56988), (13696, 48798), (4264, 16861), (76224, 8547), (1242, 58780), (51718, 93198), (92166, 83292), (17505, 57089), (75446, 7562), (72417, 44808), (16947, 12361), (13896, 62794), (72051, 10315), (32121, 12740), (25870, 55219), (911, 82802), (79173, 22178), (29644, 47610), (98329, 86306), (1573, 9702), (54652, 3749), (37143, 12856), (82486, 31928), (66059, 42463), (77789, 25283), (62224, 32230), (17085, 93477), (36146, 23354), (41720, 75725), (27819, 30571), (86323, 1233), (56834, 96694), (2553, 93314), (14486, 60769), (27299, 19241), (67047, 2143), (49020, 67359), (24527, 64111), (92412, 80476), (60828, 55554), (80388, 58241), (21380, 70807), (44904, 90337), (63599, 87177), (82483, 8223), (25684, 42863), (14794, 43828), (28246, 32114), (6236, 74233), (87466, 52368), (10049, 49396), (44481, 97613), (36967, 70944), (20194, 147), (91124, 23615), (15853, 74922), (53423, 95613), (37469, 21104), (31926, 96129), (31929, 21898), (66219, 55241), (75630, 77447), (58636, 39097), (80966, 43689), (15303, 65762), (79720, 17151), (7273, 60011), (64475, 91935), (7483, 56658), (98771, 95267), (74343, 6255), (29937, 64010), (84475, 45453), (20803, 82245), (22577, 51081), (45017, 78267), (64233, 71076), (23864, 12569), (7559, 45808), (61360, 86903), (98145, 78293), (26024, 4637), (55171, 49428), (49155, 10491), (33950, 24131), (72770, 84404), (44475, 75441), (50245, 41065), (52503, 19456), (57998, 5835), (18696, 39762), (17194, 9390), (19974, 78853), (64238, 34233), (42362, 52667), (81664, 21859), (4556, 28664), (75594, 25715), (79301, 60007), (60692, 33093), (5221, 35519), (37458, 74607), (57376, 91270), (92278, 34396), (30969, 54968), (79780, 28580), (39106, 59471), (49383, 5560), (79502, 84001), (20194, 73607), (41712, 30060), (36648, 35905), (28779, 1504), (98770, 16728), (83604, 62240), (86626, 15677), (76168, 96618), (89229, 34323), (95898, 27200), (55349, 7625), (47454, 77793), (80657, 19195), (24845, 31580), (92953, 98868), (30274, 34555), (81784, 84463), (19747, 10209), (99363, 36361), (37590, 28751), (61772, 21255), (82384, 17220), (59540, 29936), (95673, 53117), (29676, 66025), (25795, 32335), (39041, 57429), (85759, 90866), (74585, 73211), (92635, 11505), (49915, 97804), (64490, 20125), (78681, 9547), (46438, 47890), (80999, 33110), (6963, 24585), (20467, 66501), (17628, 12566), (44625, 84806), (56837, 54218), (8299, 88826), (31801, 32634), (56141, 94507), (68683, 84776), (73249, 98258), (90260, 13003), (52319, 89600), (86234, 88618), (49621, 12261), (19051, 41154), (37153, 20733), (92065, 42718), (19973, 59500), (99424, 27315), (54405, 10919), (6197, 60588), (70136, 11056), (3092, 32689), (33132, 36079), (90957, 83480), (29981, 42778), (15500, 65306), (16368, 67273), (30285, 10266), (48874, 41123), (43195, 45630), (71221, 95862), (45675, 61896), (64830, 45023), (89559, 90372), (5833, 49523), (66089, 70060), (57688, 24583), (50406, 6859), (46912, 74883), (81134, 17572), (10169, 19935), (6408, 86650), (79282, 95717), (42995, 57284), (16089, 37881), (19843, 44248), (87744, 22366), (8640, 38098), (9873, 2469), (55620, 92030), (53960, 73986), (55895, 71368), (9257, 19904), (18101, 96041), (23429, 1347), (8050, 60581), (68825, 26042), (71174, 67039), (23398, 78518), (70210, 96965), (34928, 74000), (13056, 89356), (34191, 50862), (1668, 27434), (11800, 66909), (19088, 80193), (77813, 19389), (12577, 9334), (22823, 43595), (27979, 16636), (51805, 84526), (45417, 95306), (80423, 93668), (54113, 6140), (38959, 56711), (90777, 29109), (45873, 77913), (50916, 26241), (19848, 3836), (18295, 95390), (79406, 38862), (68774, 58003), (17470, 46929), (50198, 61467), (12107, 66282), (61, 91650), (7804, 6660), (57632, 81481), (57556, 57170), (39107, 16246), (44795, 3367), (96085, 86382), (79572, 90753), (96146, 25821), (80611, 31581), (63209, 99236), (11599, 54084), (30644, 42274), (58553, 20703), (15261, 11435), (6390, 87774), (41473, 89416), (43274, 25195), (84332, 36672), (57497, 33129), (98984, 52740), (82619, 11258), (95082, 70422), (52233, 17060), (30124, 34964), (29156, 39335), (49188, 88986), (85858, 44454), (27323, 36776), (58850, 62538), (32742, 5846), (94139, 53934), (28971, 81117), (58617, 91532), (97859, 98502), (49643, 84947), (36153, 91194), (14628, 63010), (21449, 84970), (13625, 99685), (96173, 6766), (54999, 59747), (77739, 53171), (41149, 90221), (49485, 55635), (30741, 53438), (86658, 12457), (90974, 66195), (54029, 31670), (46795, 26729), (22473, 3730), (94738, 5995), (8595, 88119), (79882, 99932), (6479, 22690), (83255, 94692), (5373, 7503), (62655, 28869), (26740, 30823), (93278, 33724), (97034, 55468), (3014, 6368), (27360, 13859), (42761, 74322), (43112, 85140), (39757, 9218), (94056, 86768), (24183, 15292), (81698, 52277), (76071, 13716), (45717, 54411), (23463, 53618), (71958, 68370), (30968, 83579), (74293, 80143), (54782, 82605), (26482, 76695), (54904, 96213), (56129, 27732), (59154, 82822), (68906, 24068), (85197, 51005), (60835, 2385), (14759, 42237), (7022, 10969), (26815, 50861), (26988, 60598), (92260, 19737), (72771, 99196), (9287, 28224), (83486, 63019), (68229, 18172), (41812, 48037), (98838, 67824), (8106, 87295), (67279, 18012), (91264, 91618), (25119, 16960), (98153, 25228), (25651, 94554), (97987, 47726), (48372, 65373), (79787, 62654), (81763, 76579), (91383, 914), (89997, 57721), (57283, 43801), (96613, 70976), (57367, 42865), (19753, 21456), (55677, 73132), (16855, 61282), (35756, 36074), (29375, 88844), (33637, 27756), (51909, 82294), (77882, 44429), (19029, 44733), (85598, 8561), (36787, 74396), (88209, 66587), (4561, 44868), (10413, 64778), (86224, 53764), (54937, 26315), (75186, 13300), (91350, 46129), (87394, 30166), (94501, 67258), (5406, 39541), (54381, 30365), (91284, 33417), (23889, 70587), (40133, 19866), (26189, 72033), (73508, 26273), (40452, 87632), (73111, 2391), (14153, 56500), (38661, 88247), (18922, 15435), (25718, 72518), (42783, 34743), (72355, 71321), (67760, 74672), (88886, 72906), (2975, 53971), (56404, 43860), (26915, 43973), (76095, 25369), (53232, 26877), (20762, 73689), (98738, 14108), (40905, 7575), (56433, 12783), (4211, 23102), (35452, 8689), (82616, 45854), (90257, 58691), (50105, 82017), (95586, 67769), (39625, 85453), (86782, 56076), (68235, 41442), (83789, 24560), (85060, 92686), (86224, 76097), (18755, 8314), (89743, 73846), (83490, 79667), (22993, 95011), (12402, 16020), (48872, 49743), (154, 1526), (46346, 4373), (3188, 15535), (32400, 21795), (10850, 85080), (22126, 30742), (99491, 27362), (77499, 44953), (45990, 35021), (80128, 42532), (30194, 31471), (20180, 58795), (26397, 69871), (57833, 32197), (6791, 14870), (37913, 12711), (926, 27968), (110, 2085), (51576, 49937), (37399, 48329), (94353, 42222), (56941, 80503), (34208, 24152), (45212, 87327), (62837, 71095), (2150, 23233), (37555, 33506), (59036, 79980), (19480, 15245), (37309, 53595), (5340, 62148), (46344, 1292), (14736, 98630), (33415, 13680), (71907, 43784), (86793, 85849), (23319, 21744), (69751, 63913), (24354, 62235), (11848, 5806), (39308, 68653), (75154, 51918), (4883, 41948), (57756, 52542), (61436, 89828), (82431, 53165), (27623, 29413), (78307, 39566), (31803, 28109), (88348, 10347), (24712, 22022), (34806, 34576), (66898, 13968), (73471, 82760), (65723, 69290), (5663, 95537), (98657, 82656), (63929, 33435), (56570, 10744), (33362, 72717), (9970, 56511), (24701, 13069), (42360, 48781), (81561, 6256), (49358, 3600), (91403, 62796), (17844, 16692), (43743, 29578), (23818, 25367), (69460, 85447), (36203, 78055), (30, 70865), (15198, 85223), (68264, 72235), (5018, 77781), (84477, 33241), (59231, 87327), (89089, 98281), (18334, 54283), (33189, 89449), (74018, 96063), (39933, 78524), (74532, 73981), (11933, 59515), (71330, 4353), (58129, 76530), (90890, 44241), (76926, 43419), (50357, 6351), (39092, 64465), (9463, 81922), (47625, 38842), (57509, 76236), (42208, 16691), (37323, 29836), (21553, 36848), (29170, 62022), (77569, 86178), (7924, 26355), (37912, 92337), (53112, 72246), (31433, 54114), (54180, 27556), (94860, 51456), (16239, 3593), (76884, 93361), (90146, 95347), (96458, 22041), (55630, 29678), (49053, 97036), (65384, 57515), (68917, 3489), (2656, 54477), (65028, 37342), (24085, 50460), (90000, 81083), (42777, 81985), (24536, 63429), (55616, 9063), (83020, 9399), (96154, 98297), (6665, 67103), (59867, 96408), (47553, 77059), (70764, 97255), (9315, 75953), (94057, 96101), (87852, 84754), (86257, 14160), (72519, 91121), (87832, 17717), (38161, 7174), (59853, 12862), (69912, 72519), (29081, 11312), (91236, 10752), (27550, 58878), (51883, 3774), (52625, 60752), (37968, 1211), (23647, 83100), (98922, 12966), (90992, 34430), (55242, 20394), (43134, 40370), (65248, 7362), (8928, 81700), (31247, 95250), (47209, 55093), (3346, 31008), (54695, 2303), (8715, 99587), (66557, 38724), (21985, 58450), (82115, 42648), (90268, 3793), (86552, 29010), (83786, 52266), (70799, 87195), (3354, 20790), (84666, 93996), (19249, 22536), (13267, 53606), (92161, 7920), (42200, 50582), (31717, 88160), (2245, 85404), (11121, 54147), (91515, 14447), (36090, 34248), (96680, 40541), (52758, 34480), (43457, 43467), (26982, 41908), (32415, 53928), (82479, 83958), (38751, 98151), (34624, 1241), (19311, 36055), (14327, 40418), (58593, 21643), (3929, 86440), (13520, 5859), (37184, 78518), (52305, 42114), (6776, 81365), (85462, 7092), (55095, 60533), (51942, 7767), (83166, 58228), (21860, 55221), (85580, 57665), (94085, 76559), (58761, 43246), (12759, 62036), (42283, 42361), (60867, 62852), (17818, 30228), (43570, 83799), (2356, 55333), (27867, 13398), (9983, 85419), (48850, 4994), (48422, 45648), (6313, 49597), (43925, 20163), (83560, 50977), (5005, 66187), (12710, 44074), (34687, 87500), (88393, 47966), (14524, 56750), (4698, 29183), (3906, 51183), (95987, 49514), (75234, 20133), (43600, 68989), (24152, 18233), (46097, 793), (41482, 92133), (90677, 14958), (48882, 12940), (17166, 30215), (73267, 50393), (69846, 87503), (27321, 9562), (13804, 24914), (90436, 15300), (50493, 43678), (18957, 43821), (14637, 81056), (14869, 28605), (41424, 78811), (10455, 48311), (56222, 87012), (15377, 47542), (85416, 6883), (48400, 5457), (4960, 36969), (24316, 18674), (17551, 29001), (99783, 51530), (99029, 55203), (68165, 81063), (85027, 50594), (12360, 78742), (10302, 52093), (74713, 43258), (96235, 4238), (47810, 57488), (28029, 61904), (99288, 93880), (27546, 33316), (34057, 65663), (46959, 41396), (63538, 60552), (76506, 86918), (5875, 52563), (72631, 934), (62057, 69878), (3148, 977), (22749, 3209), (402, 47709), (42661, 89578), (7350, 88194), (29622, 58643), (90512, 60129), (16757, 18255), (54894, 80383), (86166, 15942), (76521, 66470), (70167, 59558), (39436, 78909), (49201, 56176), (6413, 3904), (37188, 55239), (47090, 93558), (60429, 24094), (79737, 26953), (23249, 55642), (74344, 41459), (51609, 41213), (81546, 32070), (9326, 28557), (87059, 94434), (84579, 35836), (51913, 8611), (70060, 20232), (37645, 42094), (99397, 61297), (75491, 32264), (97235, 67322), (833, 85256), (79146, 89949), (43084, 8179), (89149, 34148), (5354, 10692), (5072, 17411), (3175, 73039), (30784, 99797), (79620, 2118), (93174, 11554), (16044, 10086), (48452, 29011), (43868, 28199), (36999, 37058), (5300, 60302), (30105, 75668), (66000, 81946), (59073, 19023), (16157, 70479), (90302, 23785), (73541, 81556), (52357, 15696), (69643, 86481), (82007, 6064), (49690, 39146), (89098, 96196), (13683, 32160), (14547, 9433), (94042, 5867), (9237, 40695), (3034, 66920), (44596, 59128), (93584, 17355), (40745, 98789), (15017, 81728), (85626, 29866), (66825, 83054), (18813, 52566), (3413, 39628), (39324, 78867), (61629, 54043), (39814, 1335), (41386, 93371), (67026, 77258), (44224, 93916), (72788, 96888), (67341, 13039), (944, 73803), (9298, 68526), (35955, 9359), (97668, 25110), (82507, 14270), (41361, 44105), (63304, 78446), (22702, 54522), (55621, 78974), (38354, 61709), (52633, 69774), (69319, 14236), (21882, 36128), (11682, 81703), (83581, 55353), (39576, 58407), (48682, 65013), (18816, 43592), (99044, 61499), (96978, 82220), (22541, 97983), (43121, 95766), (2978, 93071), (37842, 53330), (64312, 41697), (76678, 80006), (30373, 78909), (4434, 2812), (56119, 71547), (94545, 6705), (99398, 84417), (38468, 85312), (31357, 85075), (75751, 70939), (20031, 62875), (47729, 40741), (10472, 13664), (70048, 78213), (64007, 89461), (71835, 23190), (95023, 98554), (9889, 92655), (64749, 86675), (43803, 93601), (64201, 87202), (18747, 49015), (76908, 90308), (95614, 86178), (62412, 71060), (5546, 20725), (45150, 64880), (22083, 87917), (74141, 90248), (22066, 23296), (87503, 42233), (21690, 50020), (43299, 96762), (78291, 91196), (30189, 4608), (12286, 27804), (38514, 28497), (3111, 57546), (93052, 24661), (92952, 10050), (22526, 95511), (9123, 66862), (22872, 82052), (78058, 72588), (21112, 15597), (95611, 68472)]

## BEGIN SOLUTION
class AddMulThread(threading.Thread):
    def __init__(self, queue):
        super().__init__()
        self.queue = queue
        self.results = []
    
    def run(self):
        while True:
            try:
                p = self.queue.get(timeout=1)
                self.results.append(add_mul(*p))
            except queue.Empty:
                break


def add_mul(a,b):
    sum_a_b = a + b
    mul_a_b = a * b
    return sum_a_b, mul_a_b
      

def addmultest():
    threads = []
    for _ in range(10):
        t = AddMulThread(q)
        threads.append(t)
        t.start()

    for t in threads:
        t.join()
    return threads

q = queue.Queue()
[q.put(i) for i in in_numbers]
threads = addmultest()
res = []
for t in threads:
    res.extend(t.results)
print(res)

### END SOLUTION
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution

import ast

data = ast.literal_eval(solution[0])
assert_equal(set(data), set([(112679, 2020619928), (109787, 1110548322), (97695, 2154658100), (30139, 177816450), (103759, 2259653340), (77887, 1265611386), (38826, 241723944), (82490, 1463034264), (55243, 476745456), (73562, 974422752), (8472, 17893071), (138351, 4780538730), (109814, 2605647120), (133167, 3720285560), (135921, 3819069108), (80114, 346016673), (145350, 5217623621), (111952, 2062648735), (61192, 151667152), (121073, 2559736992), (115651, 3230865948), (67165, 511230576), (124547, 3498927372), (54168, 731232656), (163133, 6575380882), (164454, 6607990368), (155291, 5551816230), (111029, 2483103280), (24793, 30363292), (193845, 9393491450), (21795, 110249534), (153885, 5759758746), (30827, 236485002), (105006, 1252831725), (143878, 5048454640), (94356, 684416084), (123835, 3641164044), (47246, 503920680), (69817, 311057622), (70740, 475637184), (142672, 5070948912), (41108, 220514395), (138720, 4056153759), (36643, 334752880), (130465, 3515194236), (96118, 1736112880), (144470, 5121933609), (81103, 920195320), (117795, 2600228304), (80105, 1504512516), (138610, 4312516224), (128747, 4116211480), (138515, 4769918916), (118240, 3446006256), (55548, 725724512), (167580, 6979084164), (44228, 339480555), (165525, 6847850516), (41817, 55138802), (54232, 600879807), (44180, 391928100), (103129, 1839306648), (107860, 2758798611), (101239, 1342263810), (110498, 1961999517), (134916, 3577578515), (156556, 6037745443), (54202, 514475977), (107486, 2418721149), (137382, 4674655557), (26761, 63486030), (112129, 3115859568), (60389, 899002170), (109767, 2767041266), (51134, 91365120), (61038, 469847105), (87656, 1913828620), (58720, 341439744), (106362, 1182699536), (170678, 7095465696), (117133, 3394148832), (138556, 4556954100), (13254, 20102729), (73576, 1343728335), (188031, 8813001050), (151715, 5685395586), (125994, 3784558520), (150305, 5408963950), (126916, 3353100000), (106957, 1649918952), (23021, 49453954), (84149, 1279840380), (139137, 4200369110), (101710, 1624114701), (120447, 3276113240), (171613, 7197774262), (135506, 4473028440), (93562, 756596361), (14741, 40778190), (149503, 5111390650), (130149, 3926494970), (81134, 1276503693), (181512, 8233258572), (112548, 3085799072), (95521, 2276384628), (77811, 1513522670), (127024, 3775529244), (131751, 3181914900), (79359, 1385991008), (117627, 1888134192), (98560, 716627775), (49044, 523218240), (76480, 187950396), (152407, 5746328256), (85412, 1797393115), (84583, 323710170), (167134, 6958980573), (100358, 2140251685), (117941, 3224598558), (118140, 3375948875), (175280, 7676492976), (92571, 1195400754), (101024, 1176387420), (97895, 1436500054), (83016, 1635023439), (121850, 2618230725), (33054, 254001104), (19626, 92194344), (63237, 864228302), (49872, 279072927), (100542, 1905070077), (141802, 4801171125), (97498, 1823979976), (74873, 863367526), (100868, 2537668867), (64772, 1031439067), (87266, 1588634173), (179435, 7982994624), (98208, 1260568575), (88152, 1827968255), (163259, 6645756264), (5862, 8456805), (162957, 6638450570), (37371, 156951278), (8537, 4600302), (171316, 7259870563), (78426, 935349605), (81992, 1581071616), (138796, 4416682179), (123012, 3391106420), (173479, 7453323588), (126573, 3845575760), (111268, 2570777755), (111958, 2704495785), (116618, 3293808277), (33438, 262673936), (51027, 599336010), (149919, 5185891350), (51949, 539452638), (86914, 1367301949), (83529, 1190730204), (82219, 1387839684), (53821, 179722788), (100576, 2092490844), (173650, 7537683816), (79195, 1424956284), (119305, 2412158164), (175655, 7571070334), (64783, 1046095812), (154067, 5862901200), (104699, 2724192960), (124824, 3879985280), (134837, 4193147682), (104107, 2541384870), (132848, 3620794615), (100508, 2319800235), (74369, 1224065610), (88348, 1933961715), (104085, 2183897300), (124186, 3760867749), (62197, 959093646), (93162, 1941144920), (37344, 70821360), (130682, 4233578160), (71597, 1222188690), (62056, 605299948), (58369, 240087948), (99610, 774021925), (141978, 4843410120), (62494, 668337408), (21125, 71895304), (84771, 651486528), (60022, 73004760), (144916, 4820014164), (175458, 7676690472), (74594, 999342945), (83008, 570522652), (117225, 3244860936), (29308, 209481867), (76690, 872585424), (82366, 743206065), (44861, 409221540), (81089, 1428515530), (83713, 75432622), (101351, 1755898794), (77254, 1411350840), (184635, 8486382674), (11275, 15261246), (58401, 204890348), (49999, 477510408), (114414, 2633613008), (108522, 2805063317), (103072, 1966739287), (94454, 2005479520), (110562, 1597054545), (59500, 844153684), (117445, 3159247000), (58390, 850454649), (87556, 106436259), (153528, 5495506796), (95867, 238230642), (75255, 880299734), (46540, 525260059), (69190, 143681721), (116379, 3301938180), (88638, 1572450497), (172888, 7436948112), (116382, 3379238712), (138629, 4681877508), (92187, 1513853660), (135241, 4056492648), (150776, 5544370023), (90706, 678257709), (68547, 1100893292), (58622, 648391432), (60360, 907092044), (80469, 462916988), (139834, 4580419488), (59445, 496380404), (142094, 4341923853), (107911, 2622586848), (20341, 2968518), (114739, 2151893260), (90775, 1187738466), (149036, 5107933299), (58573, 790745776), (128055, 3069014454), (53827, 699181242), (121460, 3658003779), (153077, 5857316610), (97733, 2292491692), (124655, 3537323574), (81065, 1006355886), (96871, 1367277720), (67284, 436460003), (156410, 5927509125), (64141, 423971814), (194038, 9409616857), (80598, 465015465), (93947, 1916267370), (129928, 3839642175), (103048, 1710942735), (73658, 1153255737), (123284, 3523345539), (135309, 4565424708), (36433, 299946616), (53367, 346262672), (148263, 5332368080), (176438, 7684066485), (30661, 120673288), (104599, 2726992188), (59646, 515685105), (58081, 819247450), (157174, 6142079080), (119916, 3355238475), (91310, 2063310925), (71959, 1021498368), (63833, 338418330), (58458, 743390352), (26584, 161451660), (98827, 1575009822), (98471, 2199059454), (95029, 2231079454), (103523, 1785093376), (33220, 130593184), (101309, 1943899710), (139308, 4758615107), (93785, 2008480356), (40740, 185444699), (112065, 2794629006), (148646, 5236707520), (126674, 3173994088), (85937, 1702303992), (108360, 2280112400), (98577, 2325672926), (54943, 274569480), (163503, 6678247502), (93801, 1486419758), (71772, 1253862720), (72553, 1315846440), (30283, 43283616), (115498, 1652224560), (145844, 5203512960), (102303, 1358035802), (172786, 7359199824), (123552, 3062606967), (123098, 2608425600), (62974, 422036125), (125247, 3691589022), (99852, 1548211115), (56425, 784605100), (191821, 9190077204), (64829, 1046118070), (166247, 6907721992), (29956, 201597123), (135724, 3612938043), (66341, 1080750090), (83027, 1312963860), (99604, 1418652480), (89476, 1782389440), (148790, 5081862741), (95701, 1959357900), (58130, 834081325), (96470, 2242085589), (176625, 7792577294), (147796, 5460442435), (104140, 1065765675), (147719, 4881886660), (84615, 1297861250), (88228, 751167507), (94328, 2223915820), (114109, 2681876890), (31548, 171185355), (86968, 1361075967), (30194, 221513448), (129431, 3784467750), (111055, 3081588466), (97125, 737166974), (64435, 1037793834), (150648, 5305717487), (153459, 5822670008), (171507, 7197300242), (103263, 1173650780), (141919, 4687782400), (174852, 7641884612), (61882, 608403081), (60205, 784024854), (57886, 770293149), (134783, 3932832670), (79473, 1188393500), (126739, 2715766560), (65324, 594048195), (66785, 375463836), (81192, 775423616), (35781, 101074388), (69211, 1195369428), (174437, 7593090360), (72759, 1282527218), (80806, 1012243000), (83641, 1101124464), (40551, 310905810), (89997, 2009845502), (88825, 1970987850), (167083, 6827387502), (107571, 2827099800), (109853, 2918841090), (179931, 8093625948), (55356, 288867659), (136149, 4630195340), (82271, 1418144104), (57265, 345734754), (121795, 3512911296), (98706, 1425686648), (30104, 202719015), (93058, 555253200), (174999, 7588635194), (100279, 2462925580), (53970, 609467409), (64091, 878013064), (110110, 1962482304), (46738, 329166720), (12342, 24376437), (147650, 5118708600), (127946, 3992284560), (127263, 3989114360), (29161, 184251328), (114142, 1738438141), (24776, 31558863), (68631, 487677050), (94867, 1792340650), (138213, 4771433786), (101916, 1837164164), (167175, 6807912650), (108928, 2584672000), (102412, 1166631936), (85053, 1739022642), (29102, 45759912), (78709, 789526200), (99281, 1530723984), (97202, 1508716257), (21911, 117393718), (66418, 994968685), (44615, 465458644), (136331, 4378869430), (140723, 4328512602), (174091, 7533061564), (60253, 332253820), (95670, 2209403849), (119886, 2642427693), (123786, 3574103049), (77157, 1336086756), (23684, 76136928), (113685, 1745160050), (118268, 3085875972), (126777, 3989098322), (64399, 819849630), (111665, 3085520466), (78389, 802476174), (91711, 5590650), (14464, 51974640), (139113, 4695912992), (114726, 3290476520), (55353, 635332322), (48162, 150824765), (182467, 8300014470), (170325, 7221397716), (121967, 2482585866), (112192, 2545775991), (162445, 6272608324), (65683, 627320316), (72918, 1295444456), (79256, 1212222759), (26696, 174509535), (94164, 560875860), (130889, 3708349768), (68469, 1090288430), (121004, 3092623104), (90626, 1904818113), (151724, 5220416160), (93877, 930124702), (165504, 6695864604), (69293, 891094980), (65088, 1053255536), (68491, 1146851260), (138174, 4377043368), (130312, 3816731532), (64099, 1004830648), (121388, 3680361300), (38588, 191409732), (148073, 5077292826), (110088, 2350040607), (150149, 5365331244), (196361, 9639307218), (134590, 4217023921), (127347, 3296936682), (77638, 921710280), (106419, 1822521530), (113310, 1358208125), (102939, 650706518), (114746, 3286025253), (130910, 4133460369), (131370, 3712503929), (105120, 2753097975), (84179, 1642737558), (99115, 1079498706), (157169, 6022023930), (85699, 1711098430), (73524, 1250783555), (26203, 83824290), (100733, 567954310), (96714, 757382805), (179814, 7982768024), (29169, 147008510), (177947, 7883582460), (12876, 40313619), (91524, 1808787195), (57563, 824207020), (127002, 3145707272), (152502, 5382281912), (9382, 19193152), (41219, 379182240), (117083, 3178083042), (128252, 3670555680), (48975, 366480026), (180824, 8161051008), (39475, 369806436), (133975, 4270926346), (89787, 1043389836), (100128, 2487507687), (77081, 1258039134), (140328, 4919768460), (114547, 2588274472), (154436, 5954063899), (137387, 4525267110), (103177, 2031036990), (151117, 5282478552), (83861, 1556569428), (141976, 4899252588), (92974, 1658429608), (136202, 4345472985), (63220, 145091475), (56996, 623375883), (17991, 77024318), (77676, 1363837715), (87586, 1635418824), (111997, 1820935620), (171967, 7218592116), (37511, 262116288), (146505, 5261204234), (86401, 1239857388), (89849, 2008523044), (166662, 6703588512), (95401, 707613270), (85291, 1211829348), (182882, 8361425152), (42079, 426018240), (123381, 2476203884), (120205, 2425404654), (145713, 4676527562), (113745, 3162222756), (142441, 4998974698), (158342, 6261328777), (92297, 83524062), (147718, 5194716837), (101084, 2509052683), (167589, 6857204288), (100232, 2459036455), (41209, 423820368), (128809, 4071770364), (78137, 1032908110), (71830, 1289861944), (118219, 2609792500), (61393, 933628572), (134203, 4271799246), (122311, 3460219378), (63762, 851224257), (94159, 732804478), (111183, 2736805652), (154796, 5873572683), (49429, 204642948), (75191, 674533314), (139988, 4635747136), (81252, 1445667155), (88486, 999973800), (137479, 4213884150), (117560, 2636327404), (161759, 6355948258), (44947, 213758646), (84746, 1651279065), (124701, 3050437428), (94476, 1686252843), (59999, 797282178), (98222, 1886472237), (99781, 1931275684), (128084, 3544889664), (75502, 174808401), (70653, 799644500), (126908, 3411717267), (34357, 292061070), (98236, 1865017924), (77526, 1486409769), (143676, 5160430955), (142432, 5059774720), (161792, 6480322716), (56946, 160563725), (100264, 2473879440), (70888, 1183533295), (101464, 1930454055), (80109, 1430716464), (94451, 1529931018), (112846, 1392995704), (48480, 309855375), (69216, 721383039), (27313, 97282522), (44141, 308042428), (128470, 3788274064), (148948, 5297273587), (132122, 4109461785), (163355, 6477767634), (125078, 3386075125), (142858, 4866387432), (109677, 2827794870), (108349, 2057857840), (177746, 7883871160), (162321, 6561387728), (27069, 155929070), (163589, 6627161578), (163157, 6651397830), (118004, 2184587923), (28422, 198680040), (98615, 2431039896), (1680, 235004), (50719, 202671058), (18723, 49525580), (54195, 706158000), (95930, 923118000), (52868, 680197492), (126853, 2722272742), (122452, 3483812547), (81011, 1610615790), (122660, 3408004096), (61665, 950235374), (78975, 1186483100), (96268, 1844384787), (90030, 1862049101), (21661, 100982170), (50624, 481912143), (28894, 25898368), (2195, 229350), (101513, 2575550712), (85728, 1807456271), (136575, 3983772366), (137444, 4583921323), (58360, 826191616), (132539, 3948228324), (133932, 4467396515), (25383, 49950950), (71061, 1258317830), (139016, 4721699280), (34725, 296972600), (90904, 1999575855), (67488, 331870320), (47636, 59876448), (113366, 1453411680), (47095, 457117200), (115691, 3148376088), (172642, 7451092257), (45063, 507048336), (133664, 4457995663), (86589, 1515671190), (17654, 68789488), (107961, 2698612124), (127072, 3901845372), (46831, 204832084), (110298, 3034615752), (151264, 5518673008), (135596, 4382444115), (57036, 812475299), (117873, 3098294762), (59912, 893950527), (98695, 914136756), (46734, 544207664), (69382, 1203452256), (80866, 934431264), (156231, 6080459960), (135013, 4553946670), (101200, 541026031), (181313, 8154592992), (97364, 2137466115), (67314, 607788080), (106079, 2425984554), (66481, 563414670), (37770, 322817369), (91141, 2066363160), (87817, 510245616), (52958, 177688800), (154199, 5739742788), (34536, 297852048), (73321, 1293830454), (49185, 604191206), (154907, 5935148620), (114258, 2825825165), (70895, 2125950), (100421, 1295219154), (140499, 4931050040), (82799, 390305058), (117718, 2808099957), (146558, 5172465537), (187370, 8755756009), (72617, 995224522), (122638, 2968722861), (170081, 7110391134), (118457, 3135698892), (148513, 5513951892), (71448, 710192495), (75683, 310499490), (134659, 4448612370), (135131, 4021064490), (120345, 3340049994), (56708, 319817307), (103557, 2520065780), (91385, 775227886), (86467, 1849850250), (133745, 4384256124), (58899, 704493728), (67159, 1113569028), (58401, 794184944), (91192, 1809181740), (163747, 6684741282), (34279, 208837020), (130249, 3500680344), (125358, 3837129552), (85547, 1700965362), (81736, 1492984080), (146316, 4881116160), (19832, 58346727), (170245, 7177967124), (185493, 8595150662), (118499, 2126030778), (85308, 1650987140), (146089, 4759906908), (122899, 3760560760), (72406, 240451413), (57133, 144690912), (102370, 2428275576), (74545, 1215329100), (171083, 7297470000), (124762, 3507072345), (87965, 1556293944), (64679, 504047808), (92419, 780304980), (194451, 9451649738), (73768, 447241495), (156275, 5771657736), (124612, 3664386627), (168019, 6882152820), (85268, 707502195), (190158, 9038971757), (172606, 7445808408), (100417, 1221399120), (163640, 6608003799), (105549, 1556119544), (45335, 273767014), (72715, 769829286), (142431, 5069948328), (40393, 328964272), (101988, 980969472), (86428, 1622088900), (55657, 195806442), (113377, 3197074000), (39179, 45979248), (106747, 1965065700), (111888, 1282622652), (125422, 3132854560), (75636, 1126605348), (83504, 1741319580), (72610, 480355776), (90628, 729417600), (126497, 2976276750), (102302, 2600885437), (34354, 103752768), (56998, 125962585), (108302, 867900705), (105281, 2577353268), (80435, 1285023250), (124763, 3502040520), (94061, 342386524), (115562, 2510873520), (136052, 4379159076), (157994, 6173318805), (24144, 69729660), (178662, 7958265336), (41785, 433795464), (66873, 711190802), (100081, 729915120), (92782, 2134560400), (119877, 2796170720), (87649, 191731980), (65268, 602168787), (105962, 1322117205), (70338, 1236010320), (137221, 3919503880), (87238, 1819095840), (86924, 1888945419), (68890, 1130761656), (86343, 1748076120), (166437, 6924771882), (136902, 3803449401), (35865, 42968384), (55366, 696258105), (54745, 579068686), (80236, 1268128299), (90369, 339622760), (19379, 79213680), (115702, 2919613312), (94419, 2202772770), (88141, 551329240), (92554, 606096504), (115628, 3335065635), (59709, 403433514), (141394, 4842589848), (77081, 1207131060), (143245, 4934970700), (170644, 7203053515), (102007, 2541178206), (74795, 791517324), (84644, 1791150163), (123719, 3825612684), (48046, 538602504), (127369, 3651122430), (57689, 130364548), (41265, 373362066), (95402, 852737877), (53844, 243956900), (94070, 2210367456), (55910, 313105861), (64088, 885659775), (134537, 4259638120), (71192, 331265935), (56784, 560180540), (122187, 3035112500), (136359, 4239858638), (71274, 824237000), (33881, 137101734), (55089, 199920798), (145501, 4752700318), (95367, 1514686122), (112589, 3007920400), (42385, 440363416), (46890, 36554921), (133615, 3821861106), (105635, 1356346566), (61822, 632533080), (47381, 518670690), (123660, 3692143931), (157349, 6111734538), (36883, 261243402), (38718, 343912856), (105736, 1383670800), (94171, 2205433254), (62778, 830714697), (95693, 1186416672), (43474, 425327745), (120235, 3264666864), (58766, 505091505), (143234, 4891988664), (62919, 731053334), (92299, 587918328), (53857, 264118800), (41929, 183366240), (42990, 454076984), (46552, 508996551), (151313, 5141817990), (154232, 5466697887), (149228, 5525659395), (135621, 4301856038), (91102, 973251120), (62395, 536662086), (117971, 3231934954), (100473, 407843930), (105298, 2748501280), (89933, 1735107216), (193168, 9321157440), (60862, 917722536), (99720, 2236284791), (88355, 1943914764), (124090, 3847352976), (163424, 6649748508), (58438, 308807625), (73565, 67837354), (131935, 4336419046), (4125, 3075596), (25958, 73001541), (48111, 19179018), (132239, 3821487058), (95544, 648225900), (88265, 1737122946), (150641, 5442396048), (35012, 305899035), (135277, 4412544402), (102108, 1373658372), (142991, 5086350870), (129725, 4179006186), (118345, 3111855324), (105377, 2763915376), (10317, 25036352), (92427, 2054227932), (140648, 4405646220), (84523, 1455976326), (106690, 2149151361), (78891, 1293620858), (115803, 3082227896), (92822, 2126961717), (113616, 2615180220), (37883, 266322582), (181493, 8221329606), (120415, 3030973044), (60524, 447022843), (90292, 1417453920), (79739, 1584628630), (160694, 6092737909), (107755, 2435641624), (164557, 6546054670), (86089, 71018248), (169095, 7119103554), (51263, 352384036), (123297, 3044260052), (16046, 57244968), (22483, 88308592), (76214, 231898825), (130581, 3072150848), (81738, 168635160), (104728, 1076532396), (26130, 161819784), (77463, 1405640972), (72067, 1237033732), (74057, 1371108942), (65602, 319600600), (105773, 2277985140), (147946, 5408436000), (78096, 1123745679), (86636, 1138729203), (114087, 2147833070), (155097, 5997709796), (68053, 821795472), (156124, 6022796283), (88071, 497290448), (88836, 1945164740), (185294, 8570871208), (45843, 440045280), (23980, 137221851), (99909, 551744414), (49932, 375899715), (69954, 203035280), (103724, 2636872288), (110939, 1624150320), (139534, 4025157805), (96745, 1227309376), (115492, 2557306116), (149879, 5550083550), (71379, 988924158), (43041, 135250364), (118191, 3101365908), (115672, 3330616047), (41149, 53151690), (134757, 3864252206), (144284, 5178294708), (138140, 4153341184), (169676, 7052283744), (80380, 878059299), (74747, 69670032), (77824, 637154748), (45314, 336502845), (122778, 2452443480), (96777, 1177374890), (85466, 1824226905), (141750, 4965945584), (77224, 1237758444), (134595, 4392612854), (100063, 2366786986), (122407, 3672414942), (83555, 986825284), (58010, 790552896), (93385, 954454446), (138934, 4626459093), (97983, 2311515432), (113695, 3164962866), (62408, 820227072), (160543, 6091106956), (179198, 7973531160), (120524, 2208634803), (138887, 4129525686), (96049, 277165438), (91172, 2018113860), (106009, 2681617464), (156684, 6134700068), (109282, 2396703057), (7246, 12468408), (127666, 4015146093), (101250, 633924225), (183815, 8390880966), (123780, 3281782016), (116432, 2667696775), (146690, 5373700189), (82906, 1259449125), (88470, 1944527189), (24136, 143089408), (148261, 5478664224), (153468, 5726130227), (95025, 1665853650), (193577, 9364896742), (102544, 916265295), (151424, 5612119575), (137404, 4100004603), (151403, 5598455602), (67762, 918884205), (167216, 6945407664), (181792, 8239823292), (133472, 4434996720), (26271, 114940850), (110030, 2929332000), (110000, 1941471111), (164389, 6691076968), (45362, 514049536), (129736, 3695514199), (71710, 1084933800), (140061, 4189697838), (169487, 7139826036), (34797, 139110912), (40090, 341599944), (67011, 1097533458), (60657, 179025606), (117713, 2294755372), (103002, 934167600), (118037, 2151480786), (75985, 609982026), (104924, 1876693344), (150646, 5666074104), (36709, 329283864), (164083, 6546676392)]))

```

## Adatmegosztás

Készíts egy programot, ami threadeket használva kiszámolja listák listái elemének az összegét. Figyelj a változók szinkronizációjára!


Írd ki a végeredményt.

<!-- grade_id: cell-cb1090853f960b9f, points: 0, grade: false, solution: true -->
```python
import threading

in_numbers = [[33, 55, 21, 55], [87, 43, 62, 34, 66, 11], [86, 41, 52, 67, 84, 11, 2], [1, 3, 5]]
my_sum = 0
lock = threading.Lock()

## BEGIN SOLUTION
def work_thread(elem):
    global my_sum
    with lock:
        my_sum += sum(elem)
        

threads = []
for i in range(len(in_numbers)):
    t = threading.Thread(target=work_thread, args=(in_numbers[i],))
    t.start()
    threads.append(t)
for t in threads:
    t.join()

print(my_sum)
### END SOLUTION
```

<!-- grade_id: cell-653f9d7aa779ca80, points: 1, grade: true, solution: false -->
```python
solution = %solution

import ast

assert_equal(int(solution[0]), 819)

```

# 2. feladatsor - Multiprocessing

## Összegek és szorzatok
A multithreading fejezethez készített feladat megoldást alakítsd át úgy, hogy multiprocessing modult használva is fusson ugyan az a megoldás.

Bónusz feladat: Mérd le mind a két megoldáshoz szükséges időt.

<!-- grade_id: cell-5c227f0b8eec859e, points: 0, grade: false, solution: true -->
```python
from multiprocessing import Process, Queue
import queue
import random
import time

in_numbers = [(90303, 22376), (11273, 98514), (64060, 33635), (22089, 8050), (72660, 31099), (23101, 54786), (31038, 7788), (25814, 56676), (10704, 44539), (56234, 17328), (4461, 4011), (67005, 71346), (34680, 75134), (39880, 93287), (96237, 39684), (75533, 4581), (64673, 80677), (88697, 23255), (58604, 2588), (93777, 27296), (68452, 47199), (8752, 58413), (81743, 42804), (28604, 25564), (90382, 72751), (94608, 69846), (99486, 55805), (31045, 79984), (23501, 1292), (97615, 96230), (7981, 13814), (64278, 89607), (14369, 16458), (91281, 13725), (60680, 83198), (7918, 86438), (75796, 48039), (30980, 16266), (4783, 65034), (63216, 7524), (75564, 67108), (6343, 34765), (96831, 41889), (17360, 19283), (92437, 38028), (72008, 24110), (62439, 82031), (13640, 67463), (29424, 88371), (50037, 30068), (91456, 47154), (69640, 59107), (64092, 74423), (52108, 66132), (21016, 34532), (77334, 90246), (9885, 34343), (81428, 84097), (40454, 1363), (15523, 38709), (31890, 12290), (22936, 80193), (66163, 41697), (85549, 15690), (88271, 22227), (98651, 36265), (87749, 68807), (41933, 12269), (32073, 75413), (62073, 75309), (24130, 2631), (61296, 50833), (26630, 33759), (70541, 39226), (1854, 49280), (9035, 52003), (41170, 46486), (52176, 6544), (12616, 93746), (71654, 99024), (52576, 64557), (84850, 53706), (1747, 11507), (39891, 33685), (88925, 99106), (84162, 67553), (49430, 76564), (90610, 59695), (89416, 37500), (18693, 88264), (20623, 2398), (64220, 19929), (94855, 44282), (81873, 19837), (41495, 78952), (72962, 98651), (56916, 78590), (8941, 84621), (3690, 11051), (52925, 96578), (47519, 82630), (59781, 21353), (92598, 88914), (47276, 65272), (49924, 45597), (38566, 39245), (47442, 79582), (99900, 31851), (53408, 25951), (98448, 19179), (7905, 90655), (33360, 15684), (2542, 73938), (83991, 68416), (47845, 37567), (4018, 80565), (78621, 88513), (69613, 30745), (43067, 74874), (48425, 69715), (89708, 85572), (77058, 15513), (13430, 87594), (17974, 79921), (50883, 32133), (93995, 27855), (20902, 12152), (11838, 7788), (43259, 19978), (43449, 6423), (75213, 25329), (85927, 55875), (72254, 25244), (60634, 14239), (52867, 48001), (36559, 28213), (61387, 25879), (81579, 97856), (83025, 15183), (54787, 33365), (77432, 85827), (2565, 3297), (82022, 80935), (32549, 4822), (7959, 578), (76859, 94457), (14671, 63755), (50976, 31016), (89383, 49413), (41710, 81302), (78348, 95131), (75920, 50653), (78533, 32735), (76695, 35263), (68611, 48007), (20824, 12614), (18330, 32697), (54150, 95769), (37603, 14346), (66287, 20627), (18237, 65292), (58492, 23727), (50244, 3577), (29398, 71178), (87772, 85878), (51556, 27639), (25796, 93509), (99769, 75886), (30627, 34156), (68592, 85475), (48315, 56384), (66320, 58504), (48654, 86183), (65022, 39085), (94555, 38293), (35913, 64595), (24590, 49779), (40005, 48343), (74945, 29140), (71823, 52363), (33931, 28266), (61702, 31460), (35340, 2004), (71330, 59352), (43502, 28095), (49934, 12122), (4453, 53916), (91115, 8495), (84990, 56988), (13696, 48798), (4264, 16861), (76224, 8547), (1242, 58780), (51718, 93198), (92166, 83292), (17505, 57089), (75446, 7562), (72417, 44808), (16947, 12361), (13896, 62794), (72051, 10315), (32121, 12740), (25870, 55219), (911, 82802), (79173, 22178), (29644, 47610), (98329, 86306), (1573, 9702), (54652, 3749), (37143, 12856), (82486, 31928), (66059, 42463), (77789, 25283), (62224, 32230), (17085, 93477), (36146, 23354), (41720, 75725), (27819, 30571), (86323, 1233), (56834, 96694), (2553, 93314), (14486, 60769), (27299, 19241), (67047, 2143), (49020, 67359), (24527, 64111), (92412, 80476), (60828, 55554), (80388, 58241), (21380, 70807), (44904, 90337), (63599, 87177), (82483, 8223), (25684, 42863), (14794, 43828), (28246, 32114), (6236, 74233), (87466, 52368), (10049, 49396), (44481, 97613), (36967, 70944), (20194, 147), (91124, 23615), (15853, 74922), (53423, 95613), (37469, 21104), (31926, 96129), (31929, 21898), (66219, 55241), (75630, 77447), (58636, 39097), (80966, 43689), (15303, 65762), (79720, 17151), (7273, 60011), (64475, 91935), (7483, 56658), (98771, 95267), (74343, 6255), (29937, 64010), (84475, 45453), (20803, 82245), (22577, 51081), (45017, 78267), (64233, 71076), (23864, 12569), (7559, 45808), (61360, 86903), (98145, 78293), (26024, 4637), (55171, 49428), (49155, 10491), (33950, 24131), (72770, 84404), (44475, 75441), (50245, 41065), (52503, 19456), (57998, 5835), (18696, 39762), (17194, 9390), (19974, 78853), (64238, 34233), (42362, 52667), (81664, 21859), (4556, 28664), (75594, 25715), (79301, 60007), (60692, 33093), (5221, 35519), (37458, 74607), (57376, 91270), (92278, 34396), (30969, 54968), (79780, 28580), (39106, 59471), (49383, 5560), (79502, 84001), (20194, 73607), (41712, 30060), (36648, 35905), (28779, 1504), (98770, 16728), (83604, 62240), (86626, 15677), (76168, 96618), (89229, 34323), (95898, 27200), (55349, 7625), (47454, 77793), (80657, 19195), (24845, 31580), (92953, 98868), (30274, 34555), (81784, 84463), (19747, 10209), (99363, 36361), (37590, 28751), (61772, 21255), (82384, 17220), (59540, 29936), (95673, 53117), (29676, 66025), (25795, 32335), (39041, 57429), (85759, 90866), (74585, 73211), (92635, 11505), (49915, 97804), (64490, 20125), (78681, 9547), (46438, 47890), (80999, 33110), (6963, 24585), (20467, 66501), (17628, 12566), (44625, 84806), (56837, 54218), (8299, 88826), (31801, 32634), (56141, 94507), (68683, 84776), (73249, 98258), (90260, 13003), (52319, 89600), (86234, 88618), (49621, 12261), (19051, 41154), (37153, 20733), (92065, 42718), (19973, 59500), (99424, 27315), (54405, 10919), (6197, 60588), (70136, 11056), (3092, 32689), (33132, 36079), (90957, 83480), (29981, 42778), (15500, 65306), (16368, 67273), (30285, 10266), (48874, 41123), (43195, 45630), (71221, 95862), (45675, 61896), (64830, 45023), (89559, 90372), (5833, 49523), (66089, 70060), (57688, 24583), (50406, 6859), (46912, 74883), (81134, 17572), (10169, 19935), (6408, 86650), (79282, 95717), (42995, 57284), (16089, 37881), (19843, 44248), (87744, 22366), (8640, 38098), (9873, 2469), (55620, 92030), (53960, 73986), (55895, 71368), (9257, 19904), (18101, 96041), (23429, 1347), (8050, 60581), (68825, 26042), (71174, 67039), (23398, 78518), (70210, 96965), (34928, 74000), (13056, 89356), (34191, 50862), (1668, 27434), (11800, 66909), (19088, 80193), (77813, 19389), (12577, 9334), (22823, 43595), (27979, 16636), (51805, 84526), (45417, 95306), (80423, 93668), (54113, 6140), (38959, 56711), (90777, 29109), (45873, 77913), (50916, 26241), (19848, 3836), (18295, 95390), (79406, 38862), (68774, 58003), (17470, 46929), (50198, 61467), (12107, 66282), (61, 91650), (7804, 6660), (57632, 81481), (57556, 57170), (39107, 16246), (44795, 3367), (96085, 86382), (79572, 90753), (96146, 25821), (80611, 31581), (63209, 99236), (11599, 54084), (30644, 42274), (58553, 20703), (15261, 11435), (6390, 87774), (41473, 89416), (43274, 25195), (84332, 36672), (57497, 33129), (98984, 52740), (82619, 11258), (95082, 70422), (52233, 17060), (30124, 34964), (29156, 39335), (49188, 88986), (85858, 44454), (27323, 36776), (58850, 62538), (32742, 5846), (94139, 53934), (28971, 81117), (58617, 91532), (97859, 98502), (49643, 84947), (36153, 91194), (14628, 63010), (21449, 84970), (13625, 99685), (96173, 6766), (54999, 59747), (77739, 53171), (41149, 90221), (49485, 55635), (30741, 53438), (86658, 12457), (90974, 66195), (54029, 31670), (46795, 26729), (22473, 3730), (94738, 5995), (8595, 88119), (79882, 99932), (6479, 22690), (83255, 94692), (5373, 7503), (62655, 28869), (26740, 30823), (93278, 33724), (97034, 55468), (3014, 6368), (27360, 13859), (42761, 74322), (43112, 85140), (39757, 9218), (94056, 86768), (24183, 15292), (81698, 52277), (76071, 13716), (45717, 54411), (23463, 53618), (71958, 68370), (30968, 83579), (74293, 80143), (54782, 82605), (26482, 76695), (54904, 96213), (56129, 27732), (59154, 82822), (68906, 24068), (85197, 51005), (60835, 2385), (14759, 42237), (7022, 10969), (26815, 50861), (26988, 60598), (92260, 19737), (72771, 99196), (9287, 28224), (83486, 63019), (68229, 18172), (41812, 48037), (98838, 67824), (8106, 87295), (67279, 18012), (91264, 91618), (25119, 16960), (98153, 25228), (25651, 94554), (97987, 47726), (48372, 65373), (79787, 62654), (81763, 76579), (91383, 914), (89997, 57721), (57283, 43801), (96613, 70976), (57367, 42865), (19753, 21456), (55677, 73132), (16855, 61282), (35756, 36074), (29375, 88844), (33637, 27756), (51909, 82294), (77882, 44429), (19029, 44733), (85598, 8561), (36787, 74396), (88209, 66587), (4561, 44868), (10413, 64778), (86224, 53764), (54937, 26315), (75186, 13300), (91350, 46129), (87394, 30166), (94501, 67258), (5406, 39541), (54381, 30365), (91284, 33417), (23889, 70587), (40133, 19866), (26189, 72033), (73508, 26273), (40452, 87632), (73111, 2391), (14153, 56500), (38661, 88247), (18922, 15435), (25718, 72518), (42783, 34743), (72355, 71321), (67760, 74672), (88886, 72906), (2975, 53971), (56404, 43860), (26915, 43973), (76095, 25369), (53232, 26877), (20762, 73689), (98738, 14108), (40905, 7575), (56433, 12783), (4211, 23102), (35452, 8689), (82616, 45854), (90257, 58691), (50105, 82017), (95586, 67769), (39625, 85453), (86782, 56076), (68235, 41442), (83789, 24560), (85060, 92686), (86224, 76097), (18755, 8314), (89743, 73846), (83490, 79667), (22993, 95011), (12402, 16020), (48872, 49743), (154, 1526), (46346, 4373), (3188, 15535), (32400, 21795), (10850, 85080), (22126, 30742), (99491, 27362), (77499, 44953), (45990, 35021), (80128, 42532), (30194, 31471), (20180, 58795), (26397, 69871), (57833, 32197), (6791, 14870), (37913, 12711), (926, 27968), (110, 2085), (51576, 49937), (37399, 48329), (94353, 42222), (56941, 80503), (34208, 24152), (45212, 87327), (62837, 71095), (2150, 23233), (37555, 33506), (59036, 79980), (19480, 15245), (37309, 53595), (5340, 62148), (46344, 1292), (14736, 98630), (33415, 13680), (71907, 43784), (86793, 85849), (23319, 21744), (69751, 63913), (24354, 62235), (11848, 5806), (39308, 68653), (75154, 51918), (4883, 41948), (57756, 52542), (61436, 89828), (82431, 53165), (27623, 29413), (78307, 39566), (31803, 28109), (88348, 10347), (24712, 22022), (34806, 34576), (66898, 13968), (73471, 82760), (65723, 69290), (5663, 95537), (98657, 82656), (63929, 33435), (56570, 10744), (33362, 72717), (9970, 56511), (24701, 13069), (42360, 48781), (81561, 6256), (49358, 3600), (91403, 62796), (17844, 16692), (43743, 29578), (23818, 25367), (69460, 85447), (36203, 78055), (30, 70865), (15198, 85223), (68264, 72235), (5018, 77781), (84477, 33241), (59231, 87327), (89089, 98281), (18334, 54283), (33189, 89449), (74018, 96063), (39933, 78524), (74532, 73981), (11933, 59515), (71330, 4353), (58129, 76530), (90890, 44241), (76926, 43419), (50357, 6351), (39092, 64465), (9463, 81922), (47625, 38842), (57509, 76236), (42208, 16691), (37323, 29836), (21553, 36848), (29170, 62022), (77569, 86178), (7924, 26355), (37912, 92337), (53112, 72246), (31433, 54114), (54180, 27556), (94860, 51456), (16239, 3593), (76884, 93361), (90146, 95347), (96458, 22041), (55630, 29678), (49053, 97036), (65384, 57515), (68917, 3489), (2656, 54477), (65028, 37342), (24085, 50460), (90000, 81083), (42777, 81985), (24536, 63429), (55616, 9063), (83020, 9399), (96154, 98297), (6665, 67103), (59867, 96408), (47553, 77059), (70764, 97255), (9315, 75953), (94057, 96101), (87852, 84754), (86257, 14160), (72519, 91121), (87832, 17717), (38161, 7174), (59853, 12862), (69912, 72519), (29081, 11312), (91236, 10752), (27550, 58878), (51883, 3774), (52625, 60752), (37968, 1211), (23647, 83100), (98922, 12966), (90992, 34430), (55242, 20394), (43134, 40370), (65248, 7362), (8928, 81700), (31247, 95250), (47209, 55093), (3346, 31008), (54695, 2303), (8715, 99587), (66557, 38724), (21985, 58450), (82115, 42648), (90268, 3793), (86552, 29010), (83786, 52266), (70799, 87195), (3354, 20790), (84666, 93996), (19249, 22536), (13267, 53606), (92161, 7920), (42200, 50582), (31717, 88160), (2245, 85404), (11121, 54147), (91515, 14447), (36090, 34248), (96680, 40541), (52758, 34480), (43457, 43467), (26982, 41908), (32415, 53928), (82479, 83958), (38751, 98151), (34624, 1241), (19311, 36055), (14327, 40418), (58593, 21643), (3929, 86440), (13520, 5859), (37184, 78518), (52305, 42114), (6776, 81365), (85462, 7092), (55095, 60533), (51942, 7767), (83166, 58228), (21860, 55221), (85580, 57665), (94085, 76559), (58761, 43246), (12759, 62036), (42283, 42361), (60867, 62852), (17818, 30228), (43570, 83799), (2356, 55333), (27867, 13398), (9983, 85419), (48850, 4994), (48422, 45648), (6313, 49597), (43925, 20163), (83560, 50977), (5005, 66187), (12710, 44074), (34687, 87500), (88393, 47966), (14524, 56750), (4698, 29183), (3906, 51183), (95987, 49514), (75234, 20133), (43600, 68989), (24152, 18233), (46097, 793), (41482, 92133), (90677, 14958), (48882, 12940), (17166, 30215), (73267, 50393), (69846, 87503), (27321, 9562), (13804, 24914), (90436, 15300), (50493, 43678), (18957, 43821), (14637, 81056), (14869, 28605), (41424, 78811), (10455, 48311), (56222, 87012), (15377, 47542), (85416, 6883), (48400, 5457), (4960, 36969), (24316, 18674), (17551, 29001), (99783, 51530), (99029, 55203), (68165, 81063), (85027, 50594), (12360, 78742), (10302, 52093), (74713, 43258), (96235, 4238), (47810, 57488), (28029, 61904), (99288, 93880), (27546, 33316), (34057, 65663), (46959, 41396), (63538, 60552), (76506, 86918), (5875, 52563), (72631, 934), (62057, 69878), (3148, 977), (22749, 3209), (402, 47709), (42661, 89578), (7350, 88194), (29622, 58643), (90512, 60129), (16757, 18255), (54894, 80383), (86166, 15942), (76521, 66470), (70167, 59558), (39436, 78909), (49201, 56176), (6413, 3904), (37188, 55239), (47090, 93558), (60429, 24094), (79737, 26953), (23249, 55642), (74344, 41459), (51609, 41213), (81546, 32070), (9326, 28557), (87059, 94434), (84579, 35836), (51913, 8611), (70060, 20232), (37645, 42094), (99397, 61297), (75491, 32264), (97235, 67322), (833, 85256), (79146, 89949), (43084, 8179), (89149, 34148), (5354, 10692), (5072, 17411), (3175, 73039), (30784, 99797), (79620, 2118), (93174, 11554), (16044, 10086), (48452, 29011), (43868, 28199), (36999, 37058), (5300, 60302), (30105, 75668), (66000, 81946), (59073, 19023), (16157, 70479), (90302, 23785), (73541, 81556), (52357, 15696), (69643, 86481), (82007, 6064), (49690, 39146), (89098, 96196), (13683, 32160), (14547, 9433), (94042, 5867), (9237, 40695), (3034, 66920), (44596, 59128), (93584, 17355), (40745, 98789), (15017, 81728), (85626, 29866), (66825, 83054), (18813, 52566), (3413, 39628), (39324, 78867), (61629, 54043), (39814, 1335), (41386, 93371), (67026, 77258), (44224, 93916), (72788, 96888), (67341, 13039), (944, 73803), (9298, 68526), (35955, 9359), (97668, 25110), (82507, 14270), (41361, 44105), (63304, 78446), (22702, 54522), (55621, 78974), (38354, 61709), (52633, 69774), (69319, 14236), (21882, 36128), (11682, 81703), (83581, 55353), (39576, 58407), (48682, 65013), (18816, 43592), (99044, 61499), (96978, 82220), (22541, 97983), (43121, 95766), (2978, 93071), (37842, 53330), (64312, 41697), (76678, 80006), (30373, 78909), (4434, 2812), (56119, 71547), (94545, 6705), (99398, 84417), (38468, 85312), (31357, 85075), (75751, 70939), (20031, 62875), (47729, 40741), (10472, 13664), (70048, 78213), (64007, 89461), (71835, 23190), (95023, 98554), (9889, 92655), (64749, 86675), (43803, 93601), (64201, 87202), (18747, 49015), (76908, 90308), (95614, 86178), (62412, 71060), (5546, 20725), (45150, 64880), (22083, 87917), (74141, 90248), (22066, 23296), (87503, 42233), (21690, 50020), (43299, 96762), (78291, 91196), (30189, 4608), (12286, 27804), (38514, 28497), (3111, 57546), (93052, 24661), (92952, 10050), (22526, 95511), (9123, 66862), (22872, 82052), (78058, 72588), (21112, 15597), (95611, 68472)]

## BEGIN SOLUTION
class AddMulProcess(Process):
    def __init__(self, queue, out_queue):
        super().__init__()
        self.queue = queue
        self.out_queue = out_queue
    
    def run(self):
        while True:
            try:
                p = self.queue.get(timeout=1)
                self.out_queue.put(add_mul(*p))
            except queue.Empty:
                break


def add_mul(a,b):
    sum_a_b = a + b
    mul_a_b = a * b
    return sum_a_b, mul_a_b
      

def addmultest():
    processes = []
    for _ in range(10):
        p = AddMulProcess(q, Queue())
        processes.append(p)
        p.start()

    for p in processes:
        p.join()
    return processes

q = queue.Queue()
[q.put(i) for i in in_numbers]
processes = addmultest()

res = []
for p in processes:
    current_results = []
    while not p.out_queue.empty():
        current_results.append(p.out_queue.get())
    res.extend(current_results)
print(res)

### END SOLUTION
```

Output:
```
[(112679, 2020619928), (109787, 1110548322), (97695, 2154658100), (30139, 177816450), (103759, 2259653340), (77887, 1265611386), (38826, 241723944), (82490, 1463034264), (55243, 476745456), (73562, 974422752), (8472, 17893071), (138351, 4780538730), (109814, 2605647120), (133167, 3720285560), (135921, 3819069108), (80114, 346016673), (145350, 5217623621), (111952, 2062648735), (61192, 151667152), (121073, 2559736992), (115651, 3230865948), (67165, 511230576), (124547, 3498927372), (54168, 731232656), (163133, 6575380882), (164454, 6607990368), (155291, 5551816230), (111029, 2483103280), (24793, 30363292), (193845, 9393491450), (21795, 110249534), (153885, 5759758746), (30827, 236485002), (105006, 1252831725), (143878, 5048454640), (94356, 684416084), (123835, 3641164044), (47246, 503920680), (69817, 311057622), (70740, 475637184), (142672, 5070948912), (41108, 220514395), (138720, 4056153759), (36643, 334752880), (130465, 3515194236), (96118, 1736112880), (144470, 5121933609), (81103, 
```

<!-- grade_id: cell-6658e55c4a75eb0e, points: 1, grade: true, solution: false -->
```python
solution = %solution

import ast

data = ast.literal_eval(solution[0])
assert_equal(set(data), set([(112679, 2020619928), (109787, 1110548322), (97695, 2154658100), (30139, 177816450), (103759, 2259653340), (77887, 1265611386), (38826, 241723944), (82490, 1463034264), (55243, 476745456), (73562, 974422752), (8472, 17893071), (138351, 4780538730), (109814, 2605647120), (133167, 3720285560), (135921, 3819069108), (80114, 346016673), (145350, 5217623621), (111952, 2062648735), (61192, 151667152), (121073, 2559736992), (115651, 3230865948), (67165, 511230576), (124547, 3498927372), (54168, 731232656), (163133, 6575380882), (164454, 6607990368), (155291, 5551816230), (111029, 2483103280), (24793, 30363292), (193845, 9393491450), (21795, 110249534), (153885, 5759758746), (30827, 236485002), (105006, 1252831725), (143878, 5048454640), (94356, 684416084), (123835, 3641164044), (47246, 503920680), (69817, 311057622), (70740, 475637184), (142672, 5070948912), (41108, 220514395), (138720, 4056153759), (36643, 334752880), (130465, 3515194236), (96118, 1736112880), (144470, 5121933609), (81103, 920195320), (117795, 2600228304), (80105, 1504512516), (138610, 4312516224), (128747, 4116211480), (138515, 4769918916), (118240, 3446006256), (55548, 725724512), (167580, 6979084164), (44228, 339480555), (165525, 6847850516), (41817, 55138802), (54232, 600879807), (44180, 391928100), (103129, 1839306648), (107860, 2758798611), (101239, 1342263810), (110498, 1961999517), (134916, 3577578515), (156556, 6037745443), (54202, 514475977), (107486, 2418721149), (137382, 4674655557), (26761, 63486030), (112129, 3115859568), (60389, 899002170), (109767, 2767041266), (51134, 91365120), (61038, 469847105), (87656, 1913828620), (58720, 341439744), (106362, 1182699536), (170678, 7095465696), (117133, 3394148832), (138556, 4556954100), (13254, 20102729), (73576, 1343728335), (188031, 8813001050), (151715, 5685395586), (125994, 3784558520), (150305, 5408963950), (126916, 3353100000), (106957, 1649918952), (23021, 49453954), (84149, 1279840380), (139137, 4200369110), (101710, 1624114701), (120447, 3276113240), (171613, 7197774262), (135506, 4473028440), (93562, 756596361), (14741, 40778190), (149503, 5111390650), (130149, 3926494970), (81134, 1276503693), (181512, 8233258572), (112548, 3085799072), (95521, 2276384628), (77811, 1513522670), (127024, 3775529244), (131751, 3181914900), (79359, 1385991008), (117627, 1888134192), (98560, 716627775), (49044, 523218240), (76480, 187950396), (152407, 5746328256), (85412, 1797393115), (84583, 323710170), (167134, 6958980573), (100358, 2140251685), (117941, 3224598558), (118140, 3375948875), (175280, 7676492976), (92571, 1195400754), (101024, 1176387420), (97895, 1436500054), (83016, 1635023439), (121850, 2618230725), (33054, 254001104), (19626, 92194344), (63237, 864228302), (49872, 279072927), (100542, 1905070077), (141802, 4801171125), (97498, 1823979976), (74873, 863367526), (100868, 2537668867), (64772, 1031439067), (87266, 1588634173), (179435, 7982994624), (98208, 1260568575), (88152, 1827968255), (163259, 6645756264), (5862, 8456805), (162957, 6638450570), (37371, 156951278), (8537, 4600302), (171316, 7259870563), (78426, 935349605), (81992, 1581071616), (138796, 4416682179), (123012, 3391106420), (173479, 7453323588), (126573, 3845575760), (111268, 2570777755), (111958, 2704495785), (116618, 3293808277), (33438, 262673936), (51027, 599336010), (149919, 5185891350), (51949, 539452638), (86914, 1367301949), (83529, 1190730204), (82219, 1387839684), (53821, 179722788), (100576, 2092490844), (173650, 7537683816), (79195, 1424956284), (119305, 2412158164), (175655, 7571070334), (64783, 1046095812), (154067, 5862901200), (104699, 2724192960), (124824, 3879985280), (134837, 4193147682), (104107, 2541384870), (132848, 3620794615), (100508, 2319800235), (74369, 1224065610), (88348, 1933961715), (104085, 2183897300), (124186, 3760867749), (62197, 959093646), (93162, 1941144920), (37344, 70821360), (130682, 4233578160), (71597, 1222188690), (62056, 605299948), (58369, 240087948), (99610, 774021925), (141978, 4843410120), (62494, 668337408), (21125, 71895304), (84771, 651486528), (60022, 73004760), (144916, 4820014164), (175458, 7676690472), (74594, 999342945), (83008, 570522652), (117225, 3244860936), (29308, 209481867), (76690, 872585424), (82366, 743206065), (44861, 409221540), (81089, 1428515530), (83713, 75432622), (101351, 1755898794), (77254, 1411350840), (184635, 8486382674), (11275, 15261246), (58401, 204890348), (49999, 477510408), (114414, 2633613008), (108522, 2805063317), (103072, 1966739287), (94454, 2005479520), (110562, 1597054545), (59500, 844153684), (117445, 3159247000), (58390, 850454649), (87556, 106436259), (153528, 5495506796), (95867, 238230642), (75255, 880299734), (46540, 525260059), (69190, 143681721), (116379, 3301938180), (88638, 1572450497), (172888, 7436948112), (116382, 3379238712), (138629, 4681877508), (92187, 1513853660), (135241, 4056492648), (150776, 5544370023), (90706, 678257709), (68547, 1100893292), (58622, 648391432), (60360, 907092044), (80469, 462916988), (139834, 4580419488), (59445, 496380404), (142094, 4341923853), (107911, 2622586848), (20341, 2968518), (114739, 2151893260), (90775, 1187738466), (149036, 5107933299), (58573, 790745776), (128055, 3069014454), (53827, 699181242), (121460, 3658003779), (153077, 5857316610), (97733, 2292491692), (124655, 3537323574), (81065, 1006355886), (96871, 1367277720), (67284, 436460003), (156410, 5927509125), (64141, 423971814), (194038, 9409616857), (80598, 465015465), (93947, 1916267370), (129928, 3839642175), (103048, 1710942735), (73658, 1153255737), (123284, 3523345539), (135309, 4565424708), (36433, 299946616), (53367, 346262672), (148263, 5332368080), (176438, 7684066485), (30661, 120673288), (104599, 2726992188), (59646, 515685105), (58081, 819247450), (157174, 6142079080), (119916, 3355238475), (91310, 2063310925), (71959, 1021498368), (63833, 338418330), (58458, 743390352), (26584, 161451660), (98827, 1575009822), (98471, 2199059454), (95029, 2231079454), (103523, 1785093376), (33220, 130593184), (101309, 1943899710), (139308, 4758615107), (93785, 2008480356), (40740, 185444699), (112065, 2794629006), (148646, 5236707520), (126674, 3173994088), (85937, 1702303992), (108360, 2280112400), (98577, 2325672926), (54943, 274569480), (163503, 6678247502), (93801, 1486419758), (71772, 1253862720), (72553, 1315846440), (30283, 43283616), (115498, 1652224560), (145844, 5203512960), (102303, 1358035802), (172786, 7359199824), (123552, 3062606967), (123098, 2608425600), (62974, 422036125), (125247, 3691589022), (99852, 1548211115), (56425, 784605100), (191821, 9190077204), (64829, 1046118070), (166247, 6907721992), (29956, 201597123), (135724, 3612938043), (66341, 1080750090), (83027, 1312963860), (99604, 1418652480), (89476, 1782389440), (148790, 5081862741), (95701, 1959357900), (58130, 834081325), (96470, 2242085589), (176625, 7792577294), (147796, 5460442435), (104140, 1065765675), (147719, 4881886660), (84615, 1297861250), (88228, 751167507), (94328, 2223915820), (114109, 2681876890), (31548, 171185355), (86968, 1361075967), (30194, 221513448), (129431, 3784467750), (111055, 3081588466), (97125, 737166974), (64435, 1037793834), (150648, 5305717487), (153459, 5822670008), (171507, 7197300242), (103263, 1173650780), (141919, 4687782400), (174852, 7641884612), (61882, 608403081), (60205, 784024854), (57886, 770293149), (134783, 3932832670), (79473, 1188393500), (126739, 2715766560), (65324, 594048195), (66785, 375463836), (81192, 775423616), (35781, 101074388), (69211, 1195369428), (174437, 7593090360), (72759, 1282527218), (80806, 1012243000), (83641, 1101124464), (40551, 310905810), (89997, 2009845502), (88825, 1970987850), (167083, 6827387502), (107571, 2827099800), (109853, 2918841090), (179931, 8093625948), (55356, 288867659), (136149, 4630195340), (82271, 1418144104), (57265, 345734754), (121795, 3512911296), (98706, 1425686648), (30104, 202719015), (93058, 555253200), (174999, 7588635194), (100279, 2462925580), (53970, 609467409), (64091, 878013064), (110110, 1962482304), (46738, 329166720), (12342, 24376437), (147650, 5118708600), (127946, 3992284560), (127263, 3989114360), (29161, 184251328), (114142, 1738438141), (24776, 31558863), (68631, 487677050), (94867, 1792340650), (138213, 4771433786), (101916, 1837164164), (167175, 6807912650), (108928, 2584672000), (102412, 1166631936), (85053, 1739022642), (29102, 45759912), (78709, 789526200), (99281, 1530723984), (97202, 1508716257), (21911, 117393718), (66418, 994968685), (44615, 465458644), (136331, 4378869430), (140723, 4328512602), (174091, 7533061564), (60253, 332253820), (95670, 2209403849), (119886, 2642427693), (123786, 3574103049), (77157, 1336086756), (23684, 76136928), (113685, 1745160050), (118268, 3085875972), (126777, 3989098322), (64399, 819849630), (111665, 3085520466), (78389, 802476174), (91711, 5590650), (14464, 51974640), (139113, 4695912992), (114726, 3290476520), (55353, 635332322), (48162, 150824765), (182467, 8300014470), (170325, 7221397716), (121967, 2482585866), (112192, 2545775991), (162445, 6272608324), (65683, 627320316), (72918, 1295444456), (79256, 1212222759), (26696, 174509535), (94164, 560875860), (130889, 3708349768), (68469, 1090288430), (121004, 3092623104), (90626, 1904818113), (151724, 5220416160), (93877, 930124702), (165504, 6695864604), (69293, 891094980), (65088, 1053255536), (68491, 1146851260), (138174, 4377043368), (130312, 3816731532), (64099, 1004830648), (121388, 3680361300), (38588, 191409732), (148073, 5077292826), (110088, 2350040607), (150149, 5365331244), (196361, 9639307218), (134590, 4217023921), (127347, 3296936682), (77638, 921710280), (106419, 1822521530), (113310, 1358208125), (102939, 650706518), (114746, 3286025253), (130910, 4133460369), (131370, 3712503929), (105120, 2753097975), (84179, 1642737558), (99115, 1079498706), (157169, 6022023930), (85699, 1711098430), (73524, 1250783555), (26203, 83824290), (100733, 567954310), (96714, 757382805), (179814, 7982768024), (29169, 147008510), (177947, 7883582460), (12876, 40313619), (91524, 1808787195), (57563, 824207020), (127002, 3145707272), (152502, 5382281912), (9382, 19193152), (41219, 379182240), (117083, 3178083042), (128252, 3670555680), (48975, 366480026), (180824, 8161051008), (39475, 369806436), (133975, 4270926346), (89787, 1043389836), (100128, 2487507687), (77081, 1258039134), (140328, 4919768460), (114547, 2588274472), (154436, 5954063899), (137387, 4525267110), (103177, 2031036990), (151117, 5282478552), (83861, 1556569428), (141976, 4899252588), (92974, 1658429608), (136202, 4345472985), (63220, 145091475), (56996, 623375883), (17991, 77024318), (77676, 1363837715), (87586, 1635418824), (111997, 1820935620), (171967, 7218592116), (37511, 262116288), (146505, 5261204234), (86401, 1239857388), (89849, 2008523044), (166662, 6703588512), (95401, 707613270), (85291, 1211829348), (182882, 8361425152), (42079, 426018240), (123381, 2476203884), (120205, 2425404654), (145713, 4676527562), (113745, 3162222756), (142441, 4998974698), (158342, 6261328777), (92297, 83524062), (147718, 5194716837), (101084, 2509052683), (167589, 6857204288), (100232, 2459036455), (41209, 423820368), (128809, 4071770364), (78137, 1032908110), (71830, 1289861944), (118219, 2609792500), (61393, 933628572), (134203, 4271799246), (122311, 3460219378), (63762, 851224257), (94159, 732804478), (111183, 2736805652), (154796, 5873572683), (49429, 204642948), (75191, 674533314), (139988, 4635747136), (81252, 1445667155), (88486, 999973800), (137479, 4213884150), (117560, 2636327404), (161759, 6355948258), (44947, 213758646), (84746, 1651279065), (124701, 3050437428), (94476, 1686252843), (59999, 797282178), (98222, 1886472237), (99781, 1931275684), (128084, 3544889664), (75502, 174808401), (70653, 799644500), (126908, 3411717267), (34357, 292061070), (98236, 1865017924), (77526, 1486409769), (143676, 5160430955), (142432, 5059774720), (161792, 6480322716), (56946, 160563725), (100264, 2473879440), (70888, 1183533295), (101464, 1930454055), (80109, 1430716464), (94451, 1529931018), (112846, 1392995704), (48480, 309855375), (69216, 721383039), (27313, 97282522), (44141, 308042428), (128470, 3788274064), (148948, 5297273587), (132122, 4109461785), (163355, 6477767634), (125078, 3386075125), (142858, 4866387432), (109677, 2827794870), (108349, 2057857840), (177746, 7883871160), (162321, 6561387728), (27069, 155929070), (163589, 6627161578), (163157, 6651397830), (118004, 2184587923), (28422, 198680040), (98615, 2431039896), (1680, 235004), (50719, 202671058), (18723, 49525580), (54195, 706158000), (95930, 923118000), (52868, 680197492), (126853, 2722272742), (122452, 3483812547), (81011, 1610615790), (122660, 3408004096), (61665, 950235374), (78975, 1186483100), (96268, 1844384787), (90030, 1862049101), (21661, 100982170), (50624, 481912143), (28894, 25898368), (2195, 229350), (101513, 2575550712), (85728, 1807456271), (136575, 3983772366), (137444, 4583921323), (58360, 826191616), (132539, 3948228324), (133932, 4467396515), (25383, 49950950), (71061, 1258317830), (139016, 4721699280), (34725, 296972600), (90904, 1999575855), (67488, 331870320), (47636, 59876448), (113366, 1453411680), (47095, 457117200), (115691, 3148376088), (172642, 7451092257), (45063, 507048336), (133664, 4457995663), (86589, 1515671190), (17654, 68789488), (107961, 2698612124), (127072, 3901845372), (46831, 204832084), (110298, 3034615752), (151264, 5518673008), (135596, 4382444115), (57036, 812475299), (117873, 3098294762), (59912, 893950527), (98695, 914136756), (46734, 544207664), (69382, 1203452256), (80866, 934431264), (156231, 6080459960), (135013, 4553946670), (101200, 541026031), (181313, 8154592992), (97364, 2137466115), (67314, 607788080), (106079, 2425984554), (66481, 563414670), (37770, 322817369), (91141, 2066363160), (87817, 510245616), (52958, 177688800), (154199, 5739742788), (34536, 297852048), (73321, 1293830454), (49185, 604191206), (154907, 5935148620), (114258, 2825825165), (70895, 2125950), (100421, 1295219154), (140499, 4931050040), (82799, 390305058), (117718, 2808099957), (146558, 5172465537), (187370, 8755756009), (72617, 995224522), (122638, 2968722861), (170081, 7110391134), (118457, 3135698892), (148513, 5513951892), (71448, 710192495), (75683, 310499490), (134659, 4448612370), (135131, 4021064490), (120345, 3340049994), (56708, 319817307), (103557, 2520065780), (91385, 775227886), (86467, 1849850250), (133745, 4384256124), (58899, 704493728), (67159, 1113569028), (58401, 794184944), (91192, 1809181740), (163747, 6684741282), (34279, 208837020), (130249, 3500680344), (125358, 3837129552), (85547, 1700965362), (81736, 1492984080), (146316, 4881116160), (19832, 58346727), (170245, 7177967124), (185493, 8595150662), (118499, 2126030778), (85308, 1650987140), (146089, 4759906908), (122899, 3760560760), (72406, 240451413), (57133, 144690912), (102370, 2428275576), (74545, 1215329100), (171083, 7297470000), (124762, 3507072345), (87965, 1556293944), (64679, 504047808), (92419, 780304980), (194451, 9451649738), (73768, 447241495), (156275, 5771657736), (124612, 3664386627), (168019, 6882152820), (85268, 707502195), (190158, 9038971757), (172606, 7445808408), (100417, 1221399120), (163640, 6608003799), (105549, 1556119544), (45335, 273767014), (72715, 769829286), (142431, 5069948328), (40393, 328964272), (101988, 980969472), (86428, 1622088900), (55657, 195806442), (113377, 3197074000), (39179, 45979248), (106747, 1965065700), (111888, 1282622652), (125422, 3132854560), (75636, 1126605348), (83504, 1741319580), (72610, 480355776), (90628, 729417600), (126497, 2976276750), (102302, 2600885437), (34354, 103752768), (56998, 125962585), (108302, 867900705), (105281, 2577353268), (80435, 1285023250), (124763, 3502040520), (94061, 342386524), (115562, 2510873520), (136052, 4379159076), (157994, 6173318805), (24144, 69729660), (178662, 7958265336), (41785, 433795464), (66873, 711190802), (100081, 729915120), (92782, 2134560400), (119877, 2796170720), (87649, 191731980), (65268, 602168787), (105962, 1322117205), (70338, 1236010320), (137221, 3919503880), (87238, 1819095840), (86924, 1888945419), (68890, 1130761656), (86343, 1748076120), (166437, 6924771882), (136902, 3803449401), (35865, 42968384), (55366, 696258105), (54745, 579068686), (80236, 1268128299), (90369, 339622760), (19379, 79213680), (115702, 2919613312), (94419, 2202772770), (88141, 551329240), (92554, 606096504), (115628, 3335065635), (59709, 403433514), (141394, 4842589848), (77081, 1207131060), (143245, 4934970700), (170644, 7203053515), (102007, 2541178206), (74795, 791517324), (84644, 1791150163), (123719, 3825612684), (48046, 538602504), (127369, 3651122430), (57689, 130364548), (41265, 373362066), (95402, 852737877), (53844, 243956900), (94070, 2210367456), (55910, 313105861), (64088, 885659775), (134537, 4259638120), (71192, 331265935), (56784, 560180540), (122187, 3035112500), (136359, 4239858638), (71274, 824237000), (33881, 137101734), (55089, 199920798), (145501, 4752700318), (95367, 1514686122), (112589, 3007920400), (42385, 440363416), (46890, 36554921), (133615, 3821861106), (105635, 1356346566), (61822, 632533080), (47381, 518670690), (123660, 3692143931), (157349, 6111734538), (36883, 261243402), (38718, 343912856), (105736, 1383670800), (94171, 2205433254), (62778, 830714697), (95693, 1186416672), (43474, 425327745), (120235, 3264666864), (58766, 505091505), (143234, 4891988664), (62919, 731053334), (92299, 587918328), (53857, 264118800), (41929, 183366240), (42990, 454076984), (46552, 508996551), (151313, 5141817990), (154232, 5466697887), (149228, 5525659395), (135621, 4301856038), (91102, 973251120), (62395, 536662086), (117971, 3231934954), (100473, 407843930), (105298, 2748501280), (89933, 1735107216), (193168, 9321157440), (60862, 917722536), (99720, 2236284791), (88355, 1943914764), (124090, 3847352976), (163424, 6649748508), (58438, 308807625), (73565, 67837354), (131935, 4336419046), (4125, 3075596), (25958, 73001541), (48111, 19179018), (132239, 3821487058), (95544, 648225900), (88265, 1737122946), (150641, 5442396048), (35012, 305899035), (135277, 4412544402), (102108, 1373658372), (142991, 5086350870), (129725, 4179006186), (118345, 3111855324), (105377, 2763915376), (10317, 25036352), (92427, 2054227932), (140648, 4405646220), (84523, 1455976326), (106690, 2149151361), (78891, 1293620858), (115803, 3082227896), (92822, 2126961717), (113616, 2615180220), (37883, 266322582), (181493, 8221329606), (120415, 3030973044), (60524, 447022843), (90292, 1417453920), (79739, 1584628630), (160694, 6092737909), (107755, 2435641624), (164557, 6546054670), (86089, 71018248), (169095, 7119103554), (51263, 352384036), (123297, 3044260052), (16046, 57244968), (22483, 88308592), (76214, 231898825), (130581, 3072150848), (81738, 168635160), (104728, 1076532396), (26130, 161819784), (77463, 1405640972), (72067, 1237033732), (74057, 1371108942), (65602, 319600600), (105773, 2277985140), (147946, 5408436000), (78096, 1123745679), (86636, 1138729203), (114087, 2147833070), (155097, 5997709796), (68053, 821795472), (156124, 6022796283), (88071, 497290448), (88836, 1945164740), (185294, 8570871208), (45843, 440045280), (23980, 137221851), (99909, 551744414), (49932, 375899715), (69954, 203035280), (103724, 2636872288), (110939, 1624150320), (139534, 4025157805), (96745, 1227309376), (115492, 2557306116), (149879, 5550083550), (71379, 988924158), (43041, 135250364), (118191, 3101365908), (115672, 3330616047), (41149, 53151690), (134757, 3864252206), (144284, 5178294708), (138140, 4153341184), (169676, 7052283744), (80380, 878059299), (74747, 69670032), (77824, 637154748), (45314, 336502845), (122778, 2452443480), (96777, 1177374890), (85466, 1824226905), (141750, 4965945584), (77224, 1237758444), (134595, 4392612854), (100063, 2366786986), (122407, 3672414942), (83555, 986825284), (58010, 790552896), (93385, 954454446), (138934, 4626459093), (97983, 2311515432), (113695, 3164962866), (62408, 820227072), (160543, 6091106956), (179198, 7973531160), (120524, 2208634803), (138887, 4129525686), (96049, 277165438), (91172, 2018113860), (106009, 2681617464), (156684, 6134700068), (109282, 2396703057), (7246, 12468408), (127666, 4015146093), (101250, 633924225), (183815, 8390880966), (123780, 3281782016), (116432, 2667696775), (146690, 5373700189), (82906, 1259449125), (88470, 1944527189), (24136, 143089408), (148261, 5478664224), (153468, 5726130227), (95025, 1665853650), (193577, 9364896742), (102544, 916265295), (151424, 5612119575), (137404, 4100004603), (151403, 5598455602), (67762, 918884205), (167216, 6945407664), (181792, 8239823292), (133472, 4434996720), (26271, 114940850), (110030, 2929332000), (110000, 1941471111), (164389, 6691076968), (45362, 514049536), (129736, 3695514199), (71710, 1084933800), (140061, 4189697838), (169487, 7139826036), (34797, 139110912), (40090, 341599944), (67011, 1097533458), (60657, 179025606), (117713, 2294755372), (103002, 934167600), (118037, 2151480786), (75985, 609982026), (104924, 1876693344), (150646, 5666074104), (36709, 329283864), (164083, 6546676392)]))

```