# 1. feladatsor - Numpy
## Áfa számolás

Készíts egy Python programot ami kiszámítja nagyon sok tétel áfa tartalmát. Oldd meg NumPy tömbök segítségével és egyetlen egy műveletben.

Használd:
```python
prices = np.array([92576.45739419,95857.75423288,71493.65061907,59923.55418634, 42067.67979619,8363.08886153,20471.9374097,35454.27570623,44773.20512158,46698.53904656])
```

Áfa (Áru Forgalmi Adó), melynek kulcsa legyen 27%.

Az Áfa tartalom meghatározására a képlet:
`netto_ar * afa_kulcs`

Az eredményt tárold el a `vat_amount` változóban.


Bonusz feladat: össze tudod kapcsolni egy numpy array-be a nettó ár és az áfa tartalom oszlopokat? (tipp np.column_stack)

Az eredményt tárold el a `result` változóban.

<!-- grade_id: cell-9897d952a40b410e, points: 0, grade: false, solution: true -->
```python
import numpy as np
np.set_printoptions(precision=2)

prices = np.array([92576.45739419,95857.75423288,71493.65061907,59923.55418634, 42067.67979619,8363.08886153,20471.9374097,35454.27570623,44773.20512158,46698.53904656])
result = None
vat_amount = None
## BEGIN SOLUTION
vat = 0.27
vat_amount = prices * vat
result = np.column_stack((prices, vat_amount))
## END SOLUTION
print(vat_amount)
print(result)
```

Output:
```
[24995.64 25881.59 19303.29 16179.36 11358.27  2258.03  5527.42  9572.65
 12088.77 12608.61]
[[92576.46 24995.64]
 [95857.75 25881.59]
 [71493.65 19303.29]
 [59923.55 16179.36]
 [42067.68 11358.27]
 [ 8363.09  2258.03]
 [20471.94  5527.42]
 [35454.28  9572.65]
 [44773.21 12088.77]
 [46698.54 12608.61]]
```

<!-- grade_id: cell-b99080c958e396ff, points: 1, grade: true, solution: false -->
```python
solution = %solution
check_data = np.array([24995.64,25881.59,19303.29,16179.36,11358.27,2258.03,5527.42,9572.65,12088.77,12608.61])
assert_true(all(np.isclose(vat_amount, check_data)))
```

# 2. feladatsor - PANDAS dataframe feladat
## Időjárás

Töltsd le a met.hu napi időjárási adatait tartalmazó adatforrást innen:

https://www.met.hu/downloads.php?id=12&file=eghajlati_adatsor_1901-2019+&no=Budapest

Adatmagyarázat: https://www.met.hu/eghajlat/magyarorszag_eghajlata/eghajlati_adatsorok/Budapest/adatok/napi_adatok/index.php       

Válaszold meg a következő kérdést Pandas modul segítségével:

Hány olyan nyári nap volt az 1920tól amikor 35 c feletti hőmérsékletet mértek? Printeld a megoldást.

<!-- grade_id: cell-9fa7d20a5c2a5066, points: 0, grade: false, solution: true -->
```python
import pandas as pd
df = pd.read_csv('./BP_d.txt', sep=';')
df.columns = ['dátum', 'közép_hőmérséklet', 'max_hőmérséklet', 'min_hőmérséklet', 'csap_összeg', 'csap_fajtája', 'napfény_tartam', 'sugárzás']

## BEGIN SOLUTION
df["dátum"] = pd.to_datetime(df["dátum"])
df['ev'] = df["dátum"].dt.year
df.ev = df.ev.fillna(1900.0)
df.ev = df.ev.astype(int)
df['ho']= df["dátum"].dt.month
df.ho = df.ho.fillna(0.0)
df.ho = df.ho.astype(int)
df['nap']= df["dátum"].dt.day
df.nap = df.nap.fillna(0.0)
df.nap = df.nap.astype(int)

meleg_napok = df[(df["ev"] > 1920) & (df["max_hőmérséklet"] > 35.0)]
meleg_napok[['dátum', 'max_hőmérséklet']]
print(meleg_napok["nap"].count())
## END SOLUTION
```

Output:
```
239
```

<!-- grade_id: cell-bd3e30a89ae3fbb3, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(int(solution[0]), 239)
```