# 1. feladatsor - HTTP kezelés
## Viccek

Készíts egy Python vicc mesélő alkalmazást ami a JokeAPI v2 szerviz-t használja.

https://sv443.net/jokeapi/v2

Az alkalmazás legyen képes lekérdezni újabb vicceket amíg a felhasználó ki nem lép belőle. Oldd meg, hogy az egyrészes és kétrészes viccek is megjeleníthetőek legyenek. A kétrészes vicceknél, a csattanót ne írd ki addig amíg a felhasználó azt nem kéri (plusz egy enter).

A feladatot nem tudod itt tesztelni. Írd meg külön, és vesd össze a megoldással.


<!-- grade_id: cell-603d3addbdcd7e68, points: 1, grade: true, solution: true -->
```python

```

```python
import requests
import json
import pprint

url = "https://sv443.net/jokeapi/v2"
path = "/joke/Programming"
params = {'blacklistFlags': 'nsfw,religious,political,racist,sexist'}

headers = { 'Accept': 'application/json' }

while True:
    key = input("Another joke? (yes/no)")
    if key == "no":
        break
        
    response = requests.request("GET", url + path, headers=headers, params=params)
    joke = response.json()
    
    if joke['type'] == 'twopart':
        print(joke['setup'] + '\n<press enter>')
        input('')
        print(joke['delivery'])
    else:
        print(joke['joke'])
```

Output:
```
Another joke? (yes/no)yes
Four engineers get into a car. The car won't start.
The Mechanical engineer says "It's a broken starter".
The Electrical engineer says "Dead battery".
The Chemical engineer says "Impurities in the gasoline".
The IT engineer says "Hey guys, I have an idea: How about we all get out of the car and get back in".
Another joke? (yes/no)yes
What are bits?
<press enter>

Tiny things left when you drop your computer down the stairs.
Another joke? (yes/no)no
```

<!-- grade_id: cell-b99080c958e396ff, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution) > 0 )
```

# 2. feladatsor - Asszinkron HTTP kliens
## Időjárás

Alakítsd át az előzőekben elkészített Python vicc mesélő alkalmazást, hogy aiohttp-t használjon.


A feladatot nem tudod itt tesztelni. Írd meg külön, és vesd össze a megoldással.

<!-- grade_id: cell-c6043aebe29386b4, points: 0, grade: true, solution: true -->
```python

```

```python
import aiohttp
import asyncio

url = "https://sv443.net/jokeapi/v2"
path = "/joke/Programming"
params = {'blacklistFlags': 'nsfw,religious,political,racist,sexist'}

headers = { 'Accept': 'application/json' }

async def fetch_joke():
    async with aiohttp.ClientSession() as session:
        async with session.get(url + path, params=params) as r:
            if r.status == 200:
                return await r.json()
            else:
                return r.status_code
            
async def main():
    while True:
        key = input("Another joke? (yes/no)")
        if key == "no":
            exit(0)
            
        joke = await fetch_joke()

        if joke['type'] == 'twopart':
            print(joke['setup'] + '\n<press enter>')
            input('')
            print(joke['delivery'])
        else:
            print(joke['joke'])
            
# asyncio.run(main())
```

Output:
```
<Task pending name='Task-1' coro=<main() running at <ipython-input-6-9420433b4728>:18>>
Another joke? (yes/no)
```