# 1. feladatsor - Tesztelés, TDD
## Kalkulátor további tesztelése


Bővítsd ki a kalkulátort, úgy, hogy az alábbi követelményeket teljesítse:

- Lehessen, 2 db szám különbségét megállapítani
- Lehessen, 2 szám szorzatát megállapítani
- Lehessen, hatványozni. A hatványozás a szorzáson alapuljon.
- Bővítsük ki a szorzást, hogy lehessen akárhány számot össze szorozni.

Megoldásodnál törekedj a TDD szabályainak követésére:

- Előbb a tesztet írd meg, minden követelményre külön-külön
- Aztán a funkcionalitást alakítsd át úgy, hogy a többi teszt lehetőleg ne romoljon el. Ha elromlik akkor folyamatosan javítsd meg az eltört teszteket.
- Majd, ha még van rá időd, jöhet a szépítgetés.



```python
class Calculator:
    def sum(self, a, b, *args):
        return a + b + sum(args)
    
```

```python
import unittest

class BasicCalcTests(unittest.TestCase):
    def setUp(self):
        self.calc = Calculator()
        
    def tearDown(self):
        self.calc = None
        
    def test_sum_of_2(self):
        calc = Calculator()
        result = calc.sum(2, 3)
        self.assertEqual(5, result)
        
    def test_sum_3_int(self):
        self.assertEqual(6, self.calc.sum(2,2,2))
```

<!-- grade_id: cell-cb623efde0efa1c3, points: 0, grade: false, solution: true -->
```python
class Calculator:
    def sum(self, a, b, *args):
        return a + b + sum(args)
    
    def sub(self, a, b):
        ## BEGIN SOLUTION
        return a-b
        ## END SOLUTION
    
    def prod(self, a, b, *args):
        ## BEGIN SOLUTION
        args = list(args)
        args.append(a)
        args.append(b)
        return reduce(lambda x,y: x*y, args)
        ## END SOLUTION
    
    def pow(self, a, b):
        ## BEGIN SOLUTION
        result = a
        for _ in range(1,b):
            result *= a
        return result
        ## END SOLUTION
```

<!-- grade_id: cell-d6227cb5c3cafb27, points: 1, grade: true, solution: false -->
```python
import unittest

class BasicCalcTests(unittest.TestCase):
    def setUp(self):
        self.calc = Calculator()
        
    def tearDown(self):
        self.calc = None
        
    def test_sum_of_2(self):
        calc = Calculator()
        result = calc.sum(2, 3)
        self.assertEqual(5, result)
        
    def test_sum_3_int(self):
        self.assertEqual(6, self.calc.sum(2,2,2))
        
    ## BEGIN SOLUTION
    def test_sub(self):
        self.assertEqual(2, self.calc.sub(4, 2))
        
    def test_prod(self):
        self.assertEquall(12, self.calc.prod(2, 6))
        
    def test_pow(self):
        self.assertEqual(9, self.calc.pow(3, 2))
        
    def test_prod_3_int(self):
        self.assertEqual(24, self.calc.prod(2, 6, 2))
        
    def test_prod_4_int(self):
        self.assertEqual(120, self.calc.prod(2, 6, 2, 5))
    ## END SOLUTION
        
```