# 1. feladatsor - tuple

## Tuple módosítás
Írj egy Python programot ami kitörli egy Tuple utolsó elemét.

<!-- grade_id: cell-a3e8d85db2cd3829, points: 0, grade: false, solution: true -->
```python
tuplex = "w", 3, "r", "s", "o", "u", "r", "c", "e"
### BEGIN SOLUTION
print(tuple(list(tuplex[:-1])))
### END SOLUTION
```

Output:
```
('w', 3, 'r', 's', 'o', 'u', 'r', 'c')
```

<!-- grade_id: cell-d8bf68236e9f6fa6, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(solution[0], "('w', 3, 'r', 's', 'o', 'u', 'r', 'c')")
```

## Lista elemek szűrése

Adott az alábbi lista

`sample_list = [87, 45, 41, 65, 94, 41, 99, 94]`

Távolítsuk el a duplikált elemeket belőle. Az így keletkezett listából készítsünk egy Tuple-t `result` néven. Írjuk ki a képernyőre.

Keressük meg a Tuple minimum és maximúm értéket. Ezeket írjuk ki a képernyőre.

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
sample_list = [87, 52, 44, 53, 54, 87, 52, 53]
result = None
### BEGIN SOLUTION
no_duplicates = []
for elem in sample_list:
    if elem not in no_duplicates:
        no_duplicates.append(elem)
result = tuple(no_duplicates)
print(result)
print(min(result))
print(max(result))
### END SOLUTION
```

Output:
```
(87, 52, 44, 53, 54)
44
87
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution) == 3)
assert_equal(result, (87, 52, 44, 53, 54))
assert_equal(solution[1].strip(), "44")
assert_equal(solution[2].strip(), "87")
```