# Jupyter notebook bevezető
Ez a munkafüzet egy interaktív felület, ahol a jelölt rubrikákba kódot írva le lehet futtatni azokat. Próbáljuk is ki:

# Fizz-buzz:
Írjuk ki a számokat 1-100 ig (az 1 és a 100 is legyen benne), mindet új sorba.
- Amikor 3-al osztható számot írnánk ki, akkor helyette írjuk, hogy: "fizz"
- Amikor 5-el osztható számot írnánk ki, akkor helyette írjuk, hogy: "buzz"
- Amikor 3-al és 5-el IS osztható számot írnánk ki, akkor helyette írjuk, hogy: "fizzbuzz"

```python
# Ez a resz csak az automatikus tesztekhez kell, futtassuk le
from io import StringIO 
import sys
import base64


def decode(target: str) -> str:
    decodedBytes = base64.b64decode(target)
    decodedStr = str(decodedBytes, "utf-8")
    return decodedStr


def encode(target: str) -> str:
    encodedBytes = base64.b64encode(target.encode("utf-8"))
    encodedStr = str(encodedBytes, "utf-8")
    return encodedStr


class Solution(list):
    def __enter__(self):
        self._stdout = sys.stdout
        sys.stdout = self._stringio = StringIO()
        return self
    def __exit__(self, *args):
        self.extend(self._stringio.getvalue().splitlines())
        del self._stringio    # free up some memory
        sys.stdout = self._stdout
```

<!-- grade_id: fizzbuzz, points: 0, grade: false, solution: true -->
```python
with Solution() as solution:
    ### BEGIN SOLUTION
    for i in range(1,101):
        if i % 3 == 0 and i % 5 == 0:
            print("fizzbuzz")
        elif i % 3 == 0:
            print("fizz")
        elif i % 5 == 0:
            print("buzz")
        else:
            print(i)
    ### END SOLUTION
print(solution)
```

Output:
```
['1', '2', 'fizz', '4', 'buzz', 'fizz', '7', '8', 'fizz', 'buzz', '11', 'fizz', '13', '14', 'fizzbuzz', '16', '17', 'fizz', '19', 'buzz', 'fizz', '22', '23', 'fizz', 'buzz', '26', 'fizz', '28', '29', 'fizzbuzz', '31', '32', 'fizz', '34', 'buzz', 'fizz', '37', '38', 'fizz', 'buzz', '41', 'fizz', '43', '44', 'fizzbuzz', '46', '47', 'fizz', '49', 'buzz', 'fizz', '52', '53', 'fizz', 'buzz', '56', 'fizz', '58', '59', 'fizzbuzz', '61', '62', 'fizz', '64', 'buzz', 'fizz', '67', '68', 'fizz', 'buzz', '71', 'fizz', '73', '74', 'fizzbuzz', '76', '77', 'fizz', '79', 'buzz', 'fizz', '82', '83', 'fizz', 'buzz', '86', 'fizz', '88', '89', 'fizzbuzz', '91', '92', 'fizz', '94', 'buzz', 'fizz', '97', '98', 'fizz', 'buzz']
```

<!-- grade_id: fizzbuzz_test, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true

assert_equal(solution[0], '1')
assert_equal(solution[1], '2')
assert_equal(solution[2], 'fizz')
assert_equal(solution[3], '4')
assert_equal(solution[4], 'buzz')
assert_equal(solution[14], 'fizzbuzz')
assert_equal(
            "\n".join(solution)
            , decode("MQoyCmZpenoKNApidXp6CmZpenoKNwo4CmZpenoKYnV6egoxMQpmaXp6CjEzCjE0CmZpenpidXp6CjE2CjE3CmZpenoKMTkKYnV6egpmaXp6CjIyCjIzCmZpenoKYnV6egoyNgpmaXp6CjI4CjI5CmZpenpidXp6CjMxCjMyCmZpenoKMzQKYnV6egpmaXp6CjM3CjM4CmZpenoKYnV6ego0MQpmaXp6CjQzCjQ0CmZpenpidXp6CjQ2CjQ3CmZpenoKNDkKYnV6egpmaXp6CjUyCjUzCmZpenoKYnV6ego1NgpmaXp6CjU4CjU5CmZpenpidXp6CjYxCjYyCmZpenoKNjQKYnV6egpmaXp6CjY3CjY4CmZpenoKYnV6ego3MQpmaXp6CjczCjc0CmZpenpidXp6Cjc2Cjc3CmZpenoKNzkKYnV6egpmaXp6CjgyCjgzCmZpenoKYnV6ego4NgpmaXp6Cjg4Cjg5CmZpenpidXp6CjkxCjkyCmZpenoKOTQKYnV6egpmaXp6Cjk3Cjk4CmZpenoKYnV6eg==")
        )
```