# 1. feladatsor - Ciklusok

## Felhasználónév és jelszó
Írj programot, amely addig kér felhasználónevet és jelszót, amíg kellően erős adatokat nem adunk meg.

Felhasználónév legalább 6 karakter, jelszó legalább nyolc karakter, nem egyezik a felhasználónévvel.)

<!-- grade_id: cell-99da8cea42168c47, points: 0, grade: false, solution: true -->
```python
username = None
password = None

### BEGIN SOLUTION
while True:
    username = input('username: ')
    password = input('password: ')

    if len(username) >= 6 and len(password) >= 8 and username != password:
        break
### END SOLUTION
```

<!-- grade_id: cell-67adc014e0d56af4, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(username is not None and password is not None)
assert_true(len(username) >= 6 and len(password) >= 8 and username != password)
```

## Szám kitalálás

Írj programot, amely 1 és 100 közt gondol egy véletlen számra. Tippet kér a felhasználótól, és kiírja, hogy
a gondolt szám nagyobb-e vagy kisebb. Ha a felhasználó eltalálja a gondolt számot, akkor a program gratulál, és a
futásának vége szakad (ezt a `Gratulalok eltalaltad` kiírásával tegye meg).

<!-- grade_id: cell-0058d833213a1604, points: 0, grade: false, solution: true -->
```python
import random

szam = None

### BEGIN SOLUTION
szam = random.randint(1, 100)

while True:
    s = int(input('na milyen szamra gondoltam? '))
    if s == szam:
        print('Gratulalok eltalaltad')
        break
    elif s > szam:
        print('Nem talalt, ennel kissebb!')
    else:
        print('Nem talalt, ennel nagyobb')
### END SOLUTION
```

<!-- grade_id: cell-9ce09057eaf32821, points: 0, grade: true, solution: false -->
```python
solution = %solution

assert_true(type(szam) == int)
assert_true(len(solution))
assert_true(solution[-1] == 'Gratulalok eltalaltad')
```