# 1. feladatsor - logikai műveletek és eldöntés

## Boolean értékek

Igaz, vagy hamis lesz a következő értékek eredménye?

* `megoldas_1`: True and not False
* `megoldas_2`: True and not True
* `megoldas_3`: True or True and False
* `megoldas_4`: not True or not False
* `megoldas_5`: True and not 0
* `megoldas_6`: 52 < 52.3
* `megoldas_7`: 1 + 52 < 52.3
* `megoldas_8`: 4 != 4.0

A megoldást ilyen formaban írd:
`megoldas_1 = True` vagy `megoldas_1 = False`

<!-- grade_id: cell-8fc684254979027c, points: 0, grade: false, solution: true -->
```python
megoldas_1 = None
megoldas_2 = None
megoldas_3 = None
megoldas_4 = None
megoldas_5 = None
megoldas_6 = None
megoldas_7 = None
megoldas_8 = None

### BEGIN SOLUTION
megoldas_1 = True
megoldas_2 = False
megoldas_3 = True
megoldas_4 = True
megoldas_5 = True
megoldas_6 = True
megoldas_7 = False
megoldas_8 = False
### END SOLUTION
```

<!-- grade_id: cell-fb154f3fad67a251, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(megoldas_1)
assert_false(megoldas_2)
assert_true(megoldas_3)
assert_true(megoldas_4)
assert_true(megoldas_5)
assert_true(megoldas_6)
assert_false(megoldas_7)
assert_false(megoldas_8)
```


## Szöveg átírása kóddá
Adottak a következő változók:
```
alma = 12
korte = 44
barack = 7
```

A következő köznyelvi kijelentéseket, és írdd át kóddá:

* `megoldas_1`: Nagyobb-e az alma és a körte értéke, mint a baracké?
* `megoldas_2`: Két barack értéke egyenlő-e egy alma értékével?
* `megoldas_3`: Igaz-e, hogy egy körte értékének a fele az 22?
* `megoldas_4`: Ha összeszorozuk (alma értéke + 2) és (barack értéke - 3)-at, akkor az eredmény nagyobb lesz-e, mint korte * 3?


<!-- grade_id: cell-f8d81787dae71956, points: 0, grade: false, solution: true -->
```python
solution = %solution

alma = 12
korte = 44
barack = 7
megoldas_1 = None
megoldas_2 = None
megoldas_3 = None
megoldas_4 = None

### BEGIN SOLUTION
megoldas_1 = alma + korte > barack
megoldas_2 = barack * 2 == alma
megoldas_3 = korte / 2 == 22
megoldas_4 = (alma + 2) * (barack - 3) > korte * 3
### END SOLUTION

```

<!-- grade_id: cell-97da3be30eab74f9, points: 1, grade: true, solution: false -->
```python
assert_true(megoldas_1)
assert_false(megoldas_2)
assert_true(megoldas_3)
assert_false(megoldas_4)
```

## Egyszerű eldöntés

Kérdd be a felhasználótól a pH értékét, és döntd el, hogy savas vagy lúgos a pH? 
Számszerűen a pH-val jellemezzük

* savas: pH < 7
* lúgos: pH > 7

Ha ez megvan, akkor írdd át úgy, hogy a semleges pH is kezelve legyen.

* semleges: pH = 7


<!-- grade_id: cell-973a1b08bbace7b7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
ph = 4
ph = int(input("ph: "))
if ph < 7:
    print("savas")
elif ph == 7:
    print("semleges")
else:
    print("lúgos")
### END SOLUTION

```

Output:
```
ph: 8
lúgos
```

<!-- grade_id: cell-479dd459e8d05531, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_true("semleges" if ph == 7 else "lúgos" if ph > 7 else "savas")
```

## Vadkamera

Egy vadkamera vezérlője a következő két értéket használja:

`light` es `temperature`

Ha a `light` kisebb, mint 0.01 lux VAGY a `temperature` fagyáspont felett van, akkor írdd ki, hogy "snap".

Viszont, ha mindkettő teljesül, akkor ne írjon ki a program semmit.

<!-- grade_id: cell-6945751c47523576, points: 0, grade: false, solution: true -->
```python
light = 0.01
temperature = 5

### BEGIN SOLUTION
if (light<0.01) or (temperature>0.0):
    if not ((light<0.01) and (temperature>0.0)):
        print("snap")      
### END SOLUTION
```

Output:
```
snap
```

<!-- grade_id: cell-8e05d38eab9e7e1d, points: 0, grade: true, solution: false -->
```python
solution = %solution

assert_equal(solution[0].strip(), "snap")
```