# 1. feladatsor - dict
## Lista átalakítása

Adott számokat tartalmazó lista.
`lista = [11, 45, 8, 11, 23, 45, 23, 45, 89]`


Írj egy olyan Python programot ami egy asszociatív tömbben eltárolja a lista elemeinek előfordulási számát.

A fenti példára a kimenet: `{11: 2, 45: 3, 8: 1, 23: 2, 89: 1}`

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
lista = [11, 45, 8, 11, 23, 45, 23, 45, 89]
### BEGIN SOLUTION
count_dict = dict()
for item in lista:
    if (item in count_dict):
        count_dict[item] += 1
    else:
        count_dict[item] = 1
print(count_dict)
### END SOLUTION
```

Output:
```
{11: 2, 45: 3, 8: 1, 23: 2, 89: 1}
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_equal(solution[0].strip(), "{11: 2, 45: 3, 8: 1, 23: 2, 89: 1}")
```

## Törlés listából
Adott egy lista és egy dict.

`roll_number = [47, 64, 69, 37, 76, 83, 95, 97]`

`sample_dict ={'Jhon':47, 'Emma':69, 'Kelly':76, 'Jason':97}`

Írj egy olyan Python programot ami végig iterál a listán és kitörli azokat az elemeket amik nem szerepelnek a dictben értékként. Írasd ki az eredményt!

A fenti példára a kimenet: `[47, 69, 76, 97]`

<!-- grade_id: cell-1bc5a980cd583d68, points: 0, grade: false, solution: true -->
```python
roll_number = [47, 64, 69, 37, 76, 83, 95, 97]
sample_dict ={'Jhon':47, 'Emma':69, 'Kelly':76, 'Jason':97}
result = None
### BEGIN SOLUTION
roll_number[:] = [item for item in roll_number if item in sample_dict.values()]
### END SOLUTION
print(roll_number)
```

Output:
```
[47, 69, 76, 97]
```

<!-- grade_id: cell-4d96a53f5f5e00fb, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(solution[0].strip(), "[47, 69, 76, 97]")
```