# 1. feladatsor - Python alapok gyakorlása
## Vándorló madarak

Adott egy lista ami madár megfigyeléseket tartalmaz időrendi sorrendben. Egy bejegyzett szám az adott madár típusa.
`madarak = [2, 3, 2, 4, 5, 1, 2, 2, 3, 4, 6, 4, 2]`


Írj egy olyan Python programot ami összesíti egy `dict`-ben a madártípusokat, és hogy hányszor fordultak elő a mérés során.

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
madarak = [2, 3, 2, 4, 5, 1, 2, 2, 3, 4, 6, 4, 2]
### BEGIN SOLUTION
counts = {}
for elem in madarak:
    counts[elem] = counts.get(elem, 0) + 1
print(counts)
### END SOLUTION
```

Output:
```
{2: 5, 3: 2, 4: 3, 5: 1, 1: 1, 6: 1}
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_equal(solution[0].strip(), "{2: 5, 3: 2, 4: 3, 5: 1, 1: 1, 6: 1}")
```

## CamelCase számoló
Készíts egy függvényt `camel_case_solver` néven, ami paraméternek egy karakterláncot kap, és megadja, hogy hány szó van összefűzve CamelCase módon a megadott stringben.

Pl:
`camel_case_solver("CamelCaseTestCase")` -> 4

<!-- grade_id: cell-1bc5a980cd583d68, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def camel_case_solver(input_str):
    my_count = 1 if not input_str[0].isupper() else 0
    for element in input_str:
        if element.isupper():
            my_count += 1
    return my_count
### END SOLUTION
```

<!-- grade_id: cell-4d96a53f5f5e00fb, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true('camel_case_solver' in locals())
assert_equal(camel_case_solver('CamelCaseTestCase'), 4)
assert_equal(camel_case_solver('alma'), 1)
assert_equal(camel_case_solver('almafa'), 1)
assert_equal(camel_case_solver('Alma'), 1)
assert_equal(camel_case_solver('AlmaFa'), 2)

```

## Doboz rajzoló

Készíts egy függvényt (`draw_box` néven), ami stringet készít, amellyel szöveget tudunk kiírni egy dobozba, pl:

```python
#########
# Hello #
#########
```

A függvénynek lehessen megadni:
* A doboz szélességét
* A doboz magasságát
* A kiírandó szöveget

A szöveget mindig helyezzük horizontális és vertikális középre. A függvény térjen vissza az eredménnyel, ne írjon ki.

Ha a megadott szélesség / magasságba nem férne be a szöveg, akkor térjen vissza `None` értékkel.

<!-- grade_id: cell-fa28ec479ae79a9b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def draw_box(width, height, text):
    if width < len(text) + 2 or height < 3:
        return None
    center_y = height // 2
    text_start = (width - len(text)) // 2
    text_current = 0
    
    ret = ""
    for y in range(height):
        for x in range(width):
            if y == 0 or y == height - 1 or x == 0 or x == width - 1:
                ret += "#"
            elif y == center_y and text_start <= x < text_start + len(text):
                ret += text[text_current]
                text_current += 1
            else:
                ret += " "
        ret += "\n"
    return ret
### END SOLUTION
print(draw_box(15, 5, "hello"))
```

Output:
```
###############
#             #
#    hello    #
#             #
###############
```

<!-- grade_id: cell-7b5025d3c59272e8, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(draw_box(1, 1, "Hello"), None)
assert_equal(draw_box(5, 15, "Hello"), None)
assert_equal(draw_box(5, 7, "Hello"), None)
assert_equal(draw_box(15, 5, "hello"), "###############\n#             #\n#    hello    #\n#             #\n###############\n")
```