# 1. feladatsor - OOP
## Személyek típusa

Írj egy Python programot ami képes személyek adatát tárolni a memóriában. A megvalósításban használj osztályt (`Person` néven).

A Person típusban el kell tudni tárolni a személy kereszt- és vezetéknevét, életkorát, címét (használj példány változókat a következő neveken:)

`fname, lname, age, address`

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
class Person:
    def __init__(self, fname, lname, age, address):
        self.fname = fname
        self.lname = lname
        self.age = age
        self.address = address
### END SOLUTION
p1 = Person('Istvan', 'Kiss', 55, '1111, Bp, Nagy Utca 8')
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_equal(p1.fname, "Istvan")
assert_equal(p1.lname, "Kiss")
assert_equal(p1.age, 55)
assert_equal(p1.address, "1111, Bp, Nagy Utca 8")
```

## Személy típus viselkedése

Az Person objektum legyen képes vissza adni a teljes nevét a személynek (használj példány metódust, legyen a neve `get_name`)


<!-- grade_id: cell-1bc5a980cd583d68, points: 0, grade: false, solution: true -->
```python
# Masold ide az elozo feladatban megirt osztalyt, es azt egeszitsd ki

### BEGIN SOLUTION
class Person:
    def __init__(self, fname, lname, age, address):
        self.fname = fname
        self.lname = lname
        self.age = age
        self.address = address
    
    def get_name(self):
        return f'{self.fname} {self.lname}'
### END SOLUTION
p1 = Person('Istvan', 'Kiss', 55, '1111, Bp, Nagy Utca 8')
p1.get_name()
```

Output:
```
'Istvan Kiss'
```

<!-- grade_id: cell-4d96a53f5f5e00fb, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(p1.get_name(), "Istvan Kiss")

```

## Személy típus kiíratás

A Person objektum legyen szépen printelhető.

<!-- grade_id: cell-fa28ec479ae79a9b, points: 0, grade: false, solution: true -->
```python
# Masold ide az elozo feladatban megirt osztalyt, es azt egeszitsd ki

### BEGIN SOLUTION
class Person:
    def __init__(self, fname, lname, age, address):
        self.fname = fname
        self.lname = lname
        self.age = age
        self.address = address
    
    def get_name(self):
        return f'{self.fname} {self.lname}'
    
    def __str__(self):
        return f'Person: {self.get_name()} {self.age} {self.address}'
    
### END SOLUTION
p1 = Person('Istvan', 'Kiss', 55, '1111, Bp, Nagy Utca 8')
p1.get_name()
print(p1)
```

Output:
```
Person: Istvan Kiss 55 1111, Bp, Nagy Utca 8
```

<!-- grade_id: cell-7b5025d3c59272e8, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(Person.__str__ is not object.__str__)
```

# 2. feladatsor - Öröklődés

Egészítsük ki a `Person` példát egy `Employee` és egy `Manager` leszármazott osztállyal.

Az `Employee` osztály egy olyan személyt valósít meg aki alkalmazott. Ennek megfelelően kell tárolnod a fizetését és a beosztását, és főnökét (`Manager`). (használj példány változókat)

`fname, lname, age, address, salary, role, manager`

Nem mindenkinek van egyből főnöke, legyen ez opcionálisan megadható. Ha valaki példányosítás után akarja megadni, legyen erre egy `set_manager` belső fuknció a `Employee` osztályon.

---


A `Manager` osztály egy olyan személyt valósít meg aki foglalkoztat `Employee`kat. Tárolj a `Manager` osztályban egy `Employee` listat ami a beosztottakat tartalmazza. Legyen egy olyan függvénye, ami kiírja az összes beosztottját (`print_employees`).

`fname, lname, age, address, employees`

Lehessen szépen kiírni az `Employee` osztály és a `Manager` osztály példányainak főbb állapotát a print() függvénnyel.

<!-- grade_id: cell-418b00a802b0b724, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Employee(Person):
    def __init__(self, fname, lname, age, address, salary, role, manager=None):
        super(Employee, self).__init__(fname, lname, age, address)
        self.salary = salary
        self.role = role
        self.manager = manager

    def set_manager(self, manager):
        self.manager = manager
        
    def __str__(self):
        return f'Employee: {self.get_name()} {self.role} {self.manager}'


class Manager(Person):
    def __init__(self, fname, lname, age, address, employees):
        super(Manager, self).__init__(fname, lname, age, address)
        self.employees = employees

    def __str__(self):
        return f'Manager: {self.get_name()} ({len(self.employees)})'

    def print_employees(self):
        for emp in self.employees:
            print(f'\t {emp}')
    
### END SOLUTION
p1 = Person('Istvan', 'Kiss', 55, '1111, Bp, Nagy Utca 8')

print(p1)

e1 = Employee('Istvan', 'Kiss', 55, '1111, Bp, Nagy Utca 8', 1000.0, 'developer')
print(e1)

m1 = Manager('Istvan', 'Kiss', 55, '1111, Bp, Nagy Utca 8', [e1])
print(m1)

e1.set_manager(m1)
m1.print_employees()
```

Output:
```
Person: Istvan Kiss 55 1111, Bp, Nagy Utca 8
Employee: Istvan Kiss developer None
Manager: Istvan Kiss (1)
	 Employee: Istvan Kiss developer Manager: Istvan Kiss (1)
```

<!-- grade_id: cell-f9590ac80e179b46, points: 0, grade: true, solution: false -->
```python
solution = %solution

assert_true(Employee.__str__ is not object.__str__)
assert_true(Manager.__str__ is not object.__str__)
assert_equal(e1.manager, m1) 
assert_true(type(m1.employees) == list) 
assert_equal(m1.employees[0], e1) 
```