# 1. feladatsor - Modulok
## Saját modul import

Írd be a következő forráskódot a `my_module.py` fájlba:
```
class CustomClass:

    def hello(self):
        print("Hello")

c = CustomClass()

print("Hejj")
```

A feladatod, hogy importáld be ezt a modult.

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
import my_module
### END SOLUTION
```

Output:
```
Hejj
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(solution[0].strip(), "Hejj")
```

## Modulból változó használat

A `my_module` ból hívd meg a `c.hello()` függvényt.


<!-- grade_id: cell-1bc5a980cd583d68, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
from my_module import c
c.hello()
### END SOLUTION
```

Output:
```
Hello
```

<!-- grade_id: cell-4d96a53f5f5e00fb, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_true("Hello" in solution)
```

## Modulból class használat

A `my_module`-ból példányosítsd a `CustomClass`-t és tárold el a my_custom_class változóba!

<!-- grade_id: cell-5f3ac85c68017818, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
from my_module import CustomClass
my_custom_class = CustomClass()
### END SOLUTION
```

<!-- grade_id: cell-e4847263acee2df0, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_true("my_custom_class" in locals())
assert_true(type(my_custom_class) == CustomClass)
```