# 1. feladatsor - Iterátorok
## Névsor

Adott egy tanulói névsor:

`nevek = ["Szandra Eszti", "Zsombor Kristóf", "Ervin Karcsi", "Gabriella Gertrúd ", "Felicitás Alexander", "Aurél Domonkos", "Vilmos Mariska", "Mónika Melánia ", "Henriett Kató", "Dominika Vida", "Felicitás Salamon", "Regina Irma", "Laci Dzsenifer ", "Sándor Judit", "Lilla Bálint", "Terézia Györgyike", "Albert Ábrahám "]`

Készíts egy iterátort `Nevsor` neven, ami ezt a névsort kapja meg, és betűrendben adja vissza a neveket.

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
nevek = ["Szandra Eszti", "Zsombor Kristóf", "Ervin Karcsi", "Gabriella Gertrúd ", "Felicitás Alexander", "Aurél Domonkos", "Vilmos Mariska", "Mónika Melánia ", "Henriett Kató", "Dominika Vida", "Felicitás Salamon", "Regina Irma", "Laci Dzsenifer ", "Sándor Judit", "Lilla Bálint", "Terézia Györgyike", "Albert Ábrahám "]

## BEGIN SOLUTION
class Nevsor:
    def __init__(self, nevek):
        self.nevek = sorted(nevek)
        self.current =  0
    def __next__(self):
        if self.current >= len(self.nevek):
            raise StopIteration 
        else:
            ret_name = self.nevek[self.current]
            self.current += 1 
            return ret_name
    def __iter__(self):
        return self
### END SOLUTION
for i in Nevsor(nevek):
    print(i)
```

Output:
```
Albert Ábrahám 
Aurél Domonkos
Dominika Vida
Ervin Karcsi
Felicitás Alexander
Felicitás Salamon
Gabriella Gertrúd 
Henriett Kató
Laci Dzsenifer 
Lilla Bálint
Mónika Melánia 
Regina Irma
Szandra Eszti
Sándor Judit
Terézia Györgyike
Vilmos Mariska
Zsombor Kristóf
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal([x for x in Nevsor(nevek)], ["Albert Ábrahám ",
"Aurél Domonkos",
"Dominika Vida",
"Ervin Karcsi",
"Felicitás Alexander",
"Felicitás Salamon",
"Gabriella Gertrúd ",
"Henriett Kató",
"Laci Dzsenifer ",
"Lilla Bálint",
"Mónika Melánia ",
"Regina Irma",
"Szandra Eszti",
"Sándor Judit",
"Terézia Györgyike",
"Vilmos Mariska",
"Zsombor Kristóf"])
```

## Négyzetszámok

Készíts egy iterátort `Squares` néven, ami végigmegy egy `limit` ig a számokon 0-tól és mindegyikre visszaadja a négyzetét.

A `limit` legyen `__init__` paraméter.

<!-- grade_id: cell-1bc5a980cd583d68, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Squares(object):
    def __init__ (self, limit=200):
        self.limit = limit 
        self.x = 0
    def __next__ (self):
        self.x += 1
        if self.x > self.limit:
            raise StopIteration 
        return (self.x-1)**2
    def __getitem__(self, idx):
        self.x = 0
        while self.x < idx:
            self.__next__()
        return self.x**2
    def __iter__(self):
        return self
        
### END SOLUTION
my_squares = Squares(limit=20)
print([x for x in my_squares])
```

Output:
```
[0, 1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121, 144, 169, 196, 225, 256, 289, 324, 361]
```

<!-- grade_id: cell-4d96a53f5f5e00fb, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_equal(solution[0], "[0, 1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121, 144, 169, 196, 225, 256, 289, 324, 361]")
```

# 2. feladatsor - Generátorok
## Hét napjai generátor

Adottak a hét napjai

`days = ["Monday","Tuesday","Wednesday","Thursday","Friday","Saturday","Sunday"]`

Írj egy generátort `day_of_week` néven, ami kigenerálja ezeket.

<!-- grade_id: cell-5f3ac85c68017818, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION

def day_of_week():
    i = 0
    days = ["Monday","Tuesday","Wednesday","Thursday","Friday","Saturday","Sunday"] 
    while True:
        if i >= len(days):
            break
        yield days[i]
        i += 1
### END SOLUTION
print([x for x in day_of_week()])
```

Output:
```
['Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday', 'Sunday']
```

<!-- grade_id: cell-e4847263acee2df0, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_equal(solution[0], "['Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday', 'Sunday']")
```