# Jupyter notebook bevezető
Ez a munkafüzet egy interaktív felület, ahol a jelölt rubrikákba kódot írva le lehet futtatni azokat. Az adott kódrészletet futtatni a felső menüben található gombokkal lehet.

A megoldásaidat az adott cellába írd, töröld a
```
    # YOUR CODE HERE
    raise NotImplementedError()
```
részt. Megoldás után az adott és a következő cellát futtatsd (ezek az automatikus tesztek). Ha nem ír semmit, akkor jó a megoldás, ha hibát kapsz akkor javítsd az megoldást és futtasd mindkét cellát újra.

```python
# Ez a cella az eredmenyek ellenorzesehez kell, csak futtasd le
from io import StringIO 
import sys

class Solution(list):
    def __enter__(self):
        self._stdout = sys.stdout
        sys.stdout = self._stringio = StringIO()
        return self
    def __exit__(self, *args):
        self.extend(self._stringio.getvalue().splitlines())
        del self._stringio
        sys.stdout = self._stdout
```

# 1. feladatsor - egyszerű utasítások 
## 1/a
A 90m<sup>2</sup>-es szobádat akarod kfiesteni, egy vödör festék 15<sup>2</sup> kifestésére elég és 5200 Ft-ba kerül.

**Mennyibe fog kerülni a teljes szoba kifestése?**

<!-- grade_id: 1_feladat, points: 0, grade: false, solution: true -->
```python
with Solution() as solution:
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
print(solution[0])
```

<!-- grade_id: 1_feladat_test, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true

assert_true(len(solution))
assert_equal(float(solution[0]),  31200)
```

## 1/b
A festő egy m<sup>2</sup> -t 12 perc alatt fest le, de az egész szobát kétszer festi át. Ha nettó 1800 forintos órabérrel dolgozik, akkor mennyi lesz a nettó végösszeg?

Mennyi lesz 27% áfával a végösszeg?

<!-- grade_id: 2_feladat, points: 0, grade: false, solution: true -->
```python
with Solution() as solution:
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
print(solution[0])
print(solution[1])
```

<!-- grade_id: 2_feladat_test, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true

assert_true(len(solution) == 2)
assert_equal(float(solution[0]),  13500)
assert_equal(float(solution[1]),  17145)
```

## 1/c
Nagymama 30kg paradicsomot akar befőzni. Kilónként 0.8 liter paradicsompüré lesz a főzés során, és ezeket 6 decis üvegekbe tölti.

**Hány üveget tud megtölteni teljesen?**

<!-- grade_id: cell-1742f85bb6e44a40, points: 0, grade: false, solution: true -->
```python
with Solution() as solution:
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
print(solution[0])
```

<!-- grade_id: cell-ca7425aa985db6c1, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true

assert_true(len(solution))
assert_equal(float(solution[0]),  40)
```

## 1/d
Az autód 7 litert fogyaszt országúton és 9-et városban. A hévízi üdülésedre 170 kilómétert utazol országúton és 35-öt városban.

**Mennyit fogyaszt az autód csak oda?**

**Mennyit fogyaszt az autód oda-vissza?**

**Mennybe kerül a teljes út, ha 350 forint a benzin?**

<!-- grade_id: cell-7a92284df926610d, points: 0, grade: false, solution: true -->
```python
with Solution() as solution:
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
print(solution[0])
print(solution[1])
print(solution[2])
```

<!-- grade_id: cell-3a685528f6ed9482, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true

assert_true(len(solution) == 3)
assert_equal(float(solution[0]),  15.05)
assert_equal(float(solution[1]),  30.1)
assert_equal(float(solution[2]),  10535.0)
```

# 2. feladatsor - változók

## 2/a

**Oldd meg az 1/d feladatot változókkal.** A szállás és étekezés a kétnapos hétvégére 35 000Ft. Összesen 70 000Ft-od van az utazásra.

**El tudsz-e menni az útra?**

<!-- grade_id: cell-09329f6eca268563, points: 0, grade: false, solution: true -->
```python
# A csak_oda, oda_vissza, teljes_ut_ara, el_tudok_e_menni nevu valtozoba szamoljuk ki az eredmenyt
csak_oda = 0
oda_vissza = 0
teljes_ut_ara = 0
el_tudok_e_menni = False

# YOUR CODE HERE
raise NotImplementedError()

print(csak_oda)
print(oda_vissza)
print(teljes_ut_ara)
print(el_tudok_e_menni)
```

<!-- grade_id: cell-6cb6b031ec2d8245, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true

assert_equal(csak_oda,  15.05)
assert_equal(oda_vissza,  30.1)
assert_equal(teljes_ut_ara,  10535.0)
assert_equal(el_tudok_e_menni,  True)
```

## 2/b

Ments el egy jelszót a "password" nevű változóban! A jelszónak legalább 8 karakter hosszúnak kell lennie, hogy biztonságosnak tekintsük. 

**Írj egy kifejezést, amely megvizsgálja, hogy biztonságos-e a megadott jelszó.**

<!-- grade_id: cell-64771b0dd0200ea6, points: 0, grade: false, solution: true -->
```python
with Solution() as solution:
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
print(solution[0])    
```

<!-- grade_id: cell-fac21acc0a8d7946, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true

assert_true('password' in locals())
assert_true(len(solution))

if len(password) > 8:
    assert_equal(solution[0],  'True')
else:
    assert_equal(solution[0],  'False')    
```

## 2/c

1 dm<sup>3</sup> egységnyi alumínium tömege 2,7kg, a nikkelé pedig 8,8kg.

**Írasd ki: nehezebb-e 0,34 dm<sup>3</sup> nikkel, mint 1,23 dm<sup>3</sup> alumínium?**

_(Ne feledd, Pythonban tizedespontot kell használnod tizedesvessző helyett)_

<!-- grade_id: cell-c984a0b1b6ad2fcc, points: 0, grade: false, solution: true -->
```python
with Solution() as solution:
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
print(solution[0])
```

<!-- grade_id: cell-27b9b29d44a50d50, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true

assert_true(len(solution))
assert_equal(solution[0],  'True')
```

## 2/d

A villanyórád 25 amperes. (Vagyis a 230V hálózati feszültséggel 25 * 230 Watt teljesítményt tud kiszolgálni)

A felkapcsolt világítást 5 db 30 wattos izzó biztosítja, a nyári meleg ellen bekapcsoltad a klímádat, ez jelenleg 1500 watton üzemel, illetve az 1200 wattos mosógéped is dolgozik.

**Írasd ki: Bekapcsolhatod-e a 2000 wattos vasalógépedet, vagy kell tőle tartanod, hogy lekapcsol a megszakító?**

<!-- grade_id: cell-2066d48477a99fc8, points: 0, grade: false, solution: true -->
```python
with Solution() as solution:
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
print(solution[0])
```

<!-- grade_id: cell-207096d8c75c12a1, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true

assert_true(len(solution))
assert_equal(solution[0],  'True')
```

# 3. feladatsor - String, operátorok, vezérlés

## 3/a

**Kérd be a felhasználó életkorát és kérdezd meg mit iszik!**

Kétféle italt ismerünk: sör és kóla
    
Ha a felhasználó 18 év alatti és sört kér, közöld vele hogy "Sajnos nem adhatok."

Ha a felhasználó 60 feletti és kólát kér, közöld vele, hogy "A koffein megemelheti a vérnyomását, sajnos nem adhatok."

Ha ismeretlen italt adott meg, akkor írd ki, hogy "Sajnos csak sört és kólát tudok adni."

Minden más esetben írd ki, hogy "Parancsoljon, itt a sör." vagy "Parancsoljon, itt a kóla".


_Ahol szükséges, ne felejtsd el konvertálni az értékeket_

<!-- grade_id: cell-0991a3ca87d6d839, points: 0, grade: false, solution: true -->
```python
eletkor = 0
ital = ""

with Solution() as solution:
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
print(solution)
```

<!-- grade_id: cell-fded93d3cf780074, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true

assert_true(len(solution))

if eletkor < 18 and ital == "sör":
    assert_equal(solution[0], "Sajnos nem adhatok.")
elif eletkor > 60 and ital == "kóla":
    assert_equal(solution[0], "A koffein megemelheti a vérnyomását, sajnos nem adhatok.")
elif ital != "sör" or ital != "kóla":
    assert_equal(solution[0], "Sajnos csak sört és kólát tudok adni.")
else:
    assert_equal(solution[0], f"Parancsoljon, itt a {ital}.")
```

## 3/b

Oldd meg a 2/a (üzemanyag fogyasztás) feladatot úgy, hogz nem előre megadott értékekkel (országúti fogyasztás, országúto megtett út, városi fogyasztás, városban megtetett út, benzin ára) dolgozol, hanem ezeket a felhasználótól kéred be.

_Ahol szükséges, ne felejtsd el konvertálni az értékeket_

<!-- grade_id: cell-0401409948bd40f0, points: 0, grade: false, solution: true -->
```python
# A kovetkezo valtozokba kerjuk be az adatokat
orszaguti_km = 0
orszaguti_fogyasztas = 0
varosi_km = 0
varosi_fogyasztas = 0
benzin_ara = 0

# A csak_oda, oda_vissza, teljes_ut_ara, el_tudok_e_menni nevu valtozoba szamoljuk ki az eredmenyt
csak_oda = 0
oda_vissza = 0
teljes_ut_ara = 0
el_tudok_e_menni = False

# YOUR CODE HERE
raise NotImplementedError()

print(csak_oda)
print(oda_vissza)
print(teljes_ut_ara)
print(el_tudok_e_menni)
```

<!-- grade_id: cell-a7de36810d308589, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true

assert_equal(csak_oda,  orszaguti_fogyasztas * orszaguti_km/100 + varosi_km/100 * varosi_fogyasztas)
assert_equal(oda_vissza,  csak_oda * 2)
assert_equal(teljes_ut_ara,  oda_vissza * 350)
assert_equal(el_tudok_e_menni,  70000 > teljes_ut_ara + 35000)
```