# Fizz-buzz:
Írjuk ki a számokat 1-100 ig (az 1 és a 100 is legyen benne), mindet új sorba.
- Amikor 3-al osztható számot írnánk ki, akkor helyette írjuk, hogy: "fizz"
- Amikor 5-el osztható számot írnánk ki, akkor helyette írjuk, hogy: "buzz"
- Amikor 3-al és 5-el IS osztható számot írnánk ki, akkor helyette írjuk, hogy: "fizzbuzz"

```python
# Ez a resz csak az automatikus tesztekhez kell, futtassuk le
from io import StringIO 
import sys
import base64


def decode(target: str) -> str:
    decodedBytes = base64.b64decode(target)
    decodedStr = str(decodedBytes, "utf-8")
    return decodedStr


def encode(target: str) -> str:
    encodedBytes = base64.b64encode(target.encode("utf-8"))
    encodedStr = str(encodedBytes, "utf-8")
    return encodedStr


class Solution(list):
    def __enter__(self):
        self._stdout = sys.stdout
        sys.stdout = self._stringio = StringIO()
        return self
    def __exit__(self, *args):
        self.extend(self._stringio.getvalue().splitlines())
        del self._stringio    # free up some memory
        sys.stdout = self._stdout
```

<!-- grade_id: cell-36c1636cf11e71bd, points: 0, grade: false, solution: true -->
```python
with Solution() as solution:
    ### BEGIN SOLUTION
    for i in range(1,101):
        if i % 3 == 0 and i % 5 == 0:
            print("fizzbuzz")
        elif i % 3 == 0:
            print("fizz")
        elif i % 5 == 0:
            print("buzz")
        else:
            print(i)
    ### END SOLUTION
print(solution)
```

<!-- grade_id: cell-e8fea7c791a65945, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true

assert_equal(solution[0], '1')
assert_equal(solution[1], '2')
assert_equal(solution[2], 'fizz')
assert_equal(solution[3], '4')
assert_equal(solution[4], 'buzz')
assert_equal(solution[14], 'fizzbuzz')
assert_equal(
            "\n".join(solution)
            , decode("MQoyCmZpenoKNApidXp6CmZpenoKNwo4CmZpenoKYnV6egoxMQpmaXp6CjEzCjE0CmZpenpidXp6CjE2CjE3CmZpenoKMTkKYnV6egpmaXp6CjIyCjIzCmZpenoKYnV6egoyNgpmaXp6CjI4CjI5CmZpenpidXp6CjMxCjMyCmZpenoKMzQKYnV6egpmaXp6CjM3CjM4CmZpenoKYnV6ego0MQpmaXp6CjQzCjQ0CmZpenpidXp6CjQ2CjQ3CmZpenoKNDkKYnV6egpmaXp6CjUyCjUzCmZpenoKYnV6ego1NgpmaXp6CjU4CjU5CmZpenpidXp6CjYxCjYyCmZpenoKNjQKYnV6egpmaXp6CjY3CjY4CmZpenoKYnV6ego3MQpmaXp6CjczCjc0CmZpenpidXp6Cjc2Cjc3CmZpenoKNzkKYnV6egpmaXp6CjgyCjgzCmZpenoKYnV6ego4NgpmaXp6Cjg4Cjg5CmZpenpidXp6CjkxCjkyCmZpenoKOTQKYnV6egpmaXp6Cjk3Cjk4CmZpenoKYnV6eg==")
        )
```